A three filament mechanistic model of musculotendon force and impedance

Author response:

The following is the authors’ response to the previous reviews.

eLife assessment

This is a valuable study that develops a new model of the way muscle responds to perturbations, synthesizing models of how it responds to small and large perturbations, both of which are used to predict how muscles function for stability but also how they can be injured, and which tend to be predicted poorly by classic Hill-type models. The evidence presented to support the model is solid, since it outperforms Hill-type models in a variety of conditions. Although the combination of phenomenological and mechanistic aspects of the model may sometimes make it challenging to interpret the output, the work will be of interest to those developing realistic models of the stability and control of movement in humans or other animals.

Reviewer #1 (Public Review):

Muscle models are important tools in the fields of biomechanics and physiology. Muscle models serve a wide variety of functions, including validating existing theories, testing new hypotheses, and predicting forces produced by humans and animals in health and disease. This paper attempts to provide an alternative to Hill-type muscle models that includes contributions of titin to force enhancement over multiple time scales. Due to the significant limitations of Hill-type models, alternative models are needed and therefore the work is important and timely.

The effort to include a role for titin in muscle models is a major strength of the methods and results. The results clearly demonstrate the weaknesses of Hill models and the advantages of incorporating titin into theoretical treatments of muscle mechanics. Another strength is to address muscle mechanics over a large range of time scales.

The authors succeed in demonstrating the need to incorporate titin in muscle models, and further show that the model accurately predicts in situ force of cat soleus (Kirsch et al. 1994; Herzog & Leonard, 2002) and rabbit posts myofibrils (Leonard et al. 2010). However, it remains unclear whether the model will be practical for use with data from different muscles or preparations. Several ad hoc modifications were described in the paper, and the degree to which the model requires parameter optimization for different muscles, preparations and experiment types remains unclear.

I think the authors should state how many parameters require fitting to the data vs the total number of model parameters. It would also be interesting for the authors to discuss challenges associated with modeling ex vivo and in vivo data sets, due to differences in means of stimulation vs. model inputs.

(1) I think the authors should state how many parameters require fitting to the data vs the total number of model parameters.

The total number of model parameters are listed in Table 1. Each parameter has, in addition, references listed for the source of data (if one exists) along with how the data were used (’C’ calculate, ’F’ fit, ’E’ estimated, or ’S’ for scaled) for the specific simulations that appear in this paper. While this is a daunting number of parameters, only a few of these parameters must be updated when modeling a new musculotendon.

Similar to a Hill-type muscle model, at least 5 parameters are needed to fit the VEXAT model to a specific musculotendon: maximum isometric force (fiso), optimal contractile element (CE) length, pennation angle, maximum shortening velocity, and tendon slack length. However, similar to a Hill model, it is only possible to use this minimal set of parameters by making use of default values for the remaining set of parameters. The defaults we have used have been extracted from mammalian muscle (see Table 1) and may not be appropriate for modeling muscle tissue that differs widely in terms of the ratio of fast/slow twitch fibers, titin isoform, temperature, and scale.

Even when these defaults are appropriate, variation is the rule for biological data rather than the exception. It will always be the case that the best fit can only be obtained by fitting more of the model’s parameters to additional data. Standard measurements of the active force-length relation, passive forcelength relation, and force-velocity relations are quite helpful to improve the accuracy of the model to a specific muscle. It is challenging to improve the fit of the model’s cross-bridge (XE) and titin models because the data required are so rare. The experiments of Kirsch et al., Prado et al, and Trombitas et´ al. are unique to our knowledge. However, if more data become available, it is relatively straight forward to update the model’s parameters using the methods described in Appendix B or the code that appears online (https://github.com/mjhmilla/Millard2023VexatMuscle).

We have modified the manuscript to make it clear that, in some circumstances, the burden of parameter identification for the VEXAT model can be as low as a Hill model:

- Section 3: last two sentences of the 2nd paragraph, found at: Page 10, column 2, lines 1-12 of MillardFranklinHerzog v3.pdf and 05 MillardFranklinHerzog v2 v3 diff.pdf

- Table 1: last two sentences of the caption, found at: Page 11 of MillardFranklinHerzog v3.pdf and 05 MillardFranklinHerzog v2 v3 diff.pdf

(2) It would also be interesting for the authors to discuss challenges associated with modeling ex vivo and in vivo data sets, due to differences in means of stimulation vs. model inputs.

All of the experiments simulated in this work are in-situ or ex-vivo. So far the main challenges of simulating any experiment have been quite consistent across both in-situ and ex-vivo datasets: there are insufficient data to fit most model parameters to a specific specimen and, instead, defaults from the literature must be used. In an ideal case, a specimen would have roughly ten extra trials collected so that the maximum isometric force, optimal fiber length, active force-length relation, passive force-length relation (upto ≈ 0_._6_f_oM), and the force-velocity relations could be identified from measurements rather than relying on literature values. Since most lab specimens are viable for a small number of trials (with the exception of cat soleus), we don’t expect this situation to change in future.

However, if data are available the fitting process is pretty straight forward for either in-situ or ex-vivo data: use a standard numerical method (for example non-linear least squares, or the bisection method) to adjust the model parameters to reduce the errors between simulation and experiment. The main difficulty, as described in the previous paragraph, is the availability of data to fit as many parameters as possible for a specific specimen. As such, the fitting process really varies from experiment to experiment and depends mainly on the richness of measurements taken from a specific specimen, and from the literature in general.

Working from in-vivo data presents an entirely different set of challenges. When working with human data, for example, it’s just not possible to directly measure muscle force with tendon buckles, and so it is never completely clear how force is distributed across the many muscles that typically actuate a joint. Further, there is also uncertainty in the boundary condition of the muscle because optical motion capture markers will move with respect to the skeleton. Video fluoroscopy offers a method of improving the accuracy of measured boundary conditions, though only for a few labs due to its great expense. A final boundary condition remains impossible to measure in any case: the geometry and forces that act at the boundaries as muscle wraps over other muscles and bones. Fitting to in-vivo data are very difficult.

While this is an interesting topic, it is tangent to our already lengthy manuscript. Since these reviews are public, we’ll leave it to the motivated reader to find this text here.

Reviewer #2 (Public Review):

This model of skeletal muscle includes springs and dampers which aim to capture the effect of crossbridge and titin stiffness during the stretch of active muscle. While both crossbridge and titin stiffness have previously been incorporated, in some form, into models, this model is the first to simultaneously include both. The authors suggest that this will allow for the prediction of muscle force in response to short-, mid- and long-range stretches. All these types of stretch are likely to be experienced by muscle during in vivo perturbations, and are known to elicit different muscle responses. Hence, it is valuable to have a single model which can predict muscle force under all these physiologically relevant conditions. In addition, this model dramatically simplifies sarcomere structure to enable this muscle model to be used in multi-muscle simulations of whole-body movement.

In order to test this model, its force predictions are compared to 3 sets of experimental data which focus on short-, mid- and long-range perturbations, and to the predictions of a Hill-type muscle model. The choice of data sets is excellent and provide a robust test of the model’s ability to predict forces over a range of length perturbations. However, I find the comparison to a Hill-type muscle model to be somewhat limiting. It is well established that Hill-type models do not have any mechanism by which they can predict the effect of active muscle stretch. Hence, that the model proposed here represents an improvement over such a model is not a surprise. Many other models, some of which are also simple enough to be incorporated into whole-body simulations, have incorporated mechanistic elements which allow for the prediction of force responses to muscle stretch. And it is not clear from the results presented here that this model would outperform such models.

The paper begins by outlining the phenomenological vs mechanistic approaches taken to muscle modelling, historically. It appears, although is not directly specified, that this model combines these approaches. A somewhat mechanistic model of the response of the crossbridges and titin to active stretch is combined with a phenomenological implementation of force-length and force-velocity relationships. This combination of approaches may be useful improving the accuracy of predictions of muscle models and whole-body simulations, which is certainly a worthy goal. However, it also may limit the insight that can be gained. For example, it does not seem that this model could reflect any effect of active titin properties on muscle shortening. In addition, it is not clear to me, either physiologically or in the model, what drives the shift from the high stiffness in short-range perturbations to the somewhat lower stiffness in mid-range perturbations.

(1) It is well established that Hill-type models do not have any mechanism by which they can predict the effect of active muscle stretch.

While many muscle physiologists are aware of the limitations of the Hill model, these limitations are not so well known among computational biomechanists. There are at least two reasons for this gap: there are few comprehensive evaluations of Hill models against several experiments, and some of the differences are quite nuanced. For example, active lengthening experiments can be replicated reasonably well using a Hill model if the lengthening is done on the ascending limb of the force length curve. Clearly the story is quite different on the descending limb as shown in Figure 9. Similarly, as Figure 8 shows, by choosing the right combination of tendon model and perturbation bandwidth it is possible to get reasonably accurate responses from the Hill model to stochastic length changes. Yet when a wide variety of perturbation bandwidths, magnitudes, and tendon models are tested it is clear that the Hill model cannot, in general, replicate the response of muscle to stochastic perturbations. For these reasons we think many of the Hill model’s drawbacks have not been clearly understood by computational biomechanists for many years now.

(2) Many other models, some of which are also simple enough to be incorporated into whole-body simulations, have incorporated mechanistic elements which allow for the prediction of force responses to muscle stretch. And it is not clear from the results presented here that this model would outperform such models.

We agree that it will be valuable to benchmark other models in the literature using the same set of experiments. Hopefully we, or perhaps others, will have the good fortune to secure research funding to continue this benchmarking work. This will, however, be quite challenging: few muscle models are accompanied by a professional-quality open-source implementation. Without such an implementation it is often impossible to reproduce published results let alone provide a fair and objective evaluation of a model.

(3) For example, it does not seem that this model could reflect any effect of active titin properties on muscle shortening.

The titin model described in the paper will provide an enhancement of force during a stretch-shortening cycle. This certainly would be an interesting next experiment to simulate in a future paper.

(4) In addition, it is not clear to me, either physiologically or in the model, what drives the shift from the high stiffness in short-range perturbations to the somewhat lower stiffness in mid-range perturbations.

We can only respond to what drives the frequency dependent stiffness in the model, though we’re quite interested in what happens physiologically. Hopefully that there are some new experiments done to examine this phenomena in the future. In the case of the model, the reasons are pretty straight forward: the formulation of Eqn. 16 is responsible for this shift.

Equation 16 has been formulated so that the acceleration of the attachment point of the XE is driven by the force difference between the XE and a reference Hill model (numerator of the first term in Eqn. 16) which is then low pass filtered (denominator of the first term in Eqn. 16). Due to this formulation the attachment point moves less when the numerator is small, or when the differences in the numerator change rapidly and effectively become filtered out. When the attachment point moves less, more of the CE’s force output is determined by variations in the length of the XE and its stiffness.

On the other hand, the attachment point will move when the numerator of the first term in Eqn. 16 is large, or when those differences are not short lived. When the attachment point moves to reduce the strain in the XE, the force produced by the XE’s spring-damper is reduced. As a result, the CE’s force output is less influenced by variations of the length of the XE and its stiffness.

Reviewer #2 (Recommendations for the Authors):

I find the clarity of the manuscript to be much improved following revision. While I still find the combination of phenomenological and mechanistic approaches to be a little limiting with regards to our understanding of muscle contraction, the revised description of small length changes makes the interpretation much less confusing.

Similarly, while I agree that Hill-type models are widely used their limitations have been addressed extensively and are very well established. Hence, moving forward I think it would be much more valuable to start to compare these newer models to one another rather than just showing an improvement over a Hill model under (very biologically important) conditions which that model has no capacity to predict forces.

(1) While I still find the combination of phenomenological and mechanistic approaches to be a little limiting with regards to our understanding of muscle contraction ...

We have had to abstract some of the details of reality to have a model that can be used to simulate hundreds of muscles. In contrast, FiberSim produced by Kenneth Campbell’s group uses much less abstraction and might be of greater interest to you. FiberSim’s models include individual cross-bridges, titin molecules, and an explicit representation of the spatial geometry of a sarcomere. While this model is a great tool for testing muscle physiology questions through simulation, it is computationally expensive to use this model to simulate hundreds of muscles simultaneously.

Kosta S, Colli D, Ye Q, Campbell KS. FiberSim: A flexible open-source model of myofilament-level contraction. Biophysical journal. 2022 Jan 18;121(2):175-82.https://campbell-muscle-lab.github.io/FiberSim/

(2) Similarly, while I agree that Hill-type models are widely used their limitations have been addressed extensively and are very well established.

Please see our response 1 to Reviewer # 1.

(3) Hence, moving forward I think it would be much more valuable to start to compare these newer models to one another rather than just showing an improvement over a Hill model under (very biologically important) conditions which that model has no capacity to predict forces.

Please see our response to 2 to Reviewer #1.

Changes from the earlier version to the revised version

Abstract

Abstract

The force developed by actively lengthened muscle depends on different structures across different scales of lengthening. For small perturbations, the active response of muscle is well captured by a linear-time-invariant (LTI) system: a stiff spring in parallel with a light damper. The force response of muscle to longer stretches is better represented by a compliant spring that can fix its end when activated. Experimental work has shown that the stiffness and damping (impedance) of muscle in response to small perturbations is of fundamental importance to motor learning and mechanical stability, while the huge forces developed during long active stretches are critical for simulating and predicting injury. Outside of motor learning and injury, muscle is actively lengthened as a part of nearly all terrestrial locomotion. Despite the functional importance of impedance and active lengthening, no single muscle model has all of these mechanical properties. In this work, we present the viscoelastic-crossbridge active-titin (VEXAT) model that can replicate the response of muscle to length changes great and small. To evaluate the VEXAT model, we compare its response to biological muscle by simulating experiments that measure the impedance of muscle, and the forces developed during long active stretches. In addition, we have also compared the responses of the VEXAT model to a popular Hill-type muscle model. The VEXAT model more accurately captures the impedance of biological muscle and its responses to long active stretches than a Hill-type model and can still reproduce the force-velocity and force-length relations of muscle. While the comparison between the VEXAT model and biological muscle is favorable, there are some phenomena that can be improved: the low frequency phase response of the model, and a mechanism to support passive force enhancement.

Introduction

The stiffness and damping of muscle are properties of fundamental importance for motor control, and the accurate simulation of muscle force. The central nervous system (CNS) exploits the activation-dependent stiffness and damping (impedance) of muscle when learning new movements [1], and when moving in unstable [2] or noisy environments [3]. Reaching experiments using haptic manipulanda show that the CNS uses co-contraction to increase the stiffness of the arm when perturbed by an unstable force field [4]. With time and repetition, the force field becomes learned and co-contraction is reduced [1].

The force response of muscle is not uniform, but varies with both the length and time of perturbation. Under constant activation and at a consistent nominal length, Kirsch et al. [5] were able to show that muscle behaves like a linear-time-invariant (LTI) system in response to small1 perturbations: a spring-damper of best fit captured over 90% of the observed variation in muscle force for small perturbations (1-3.8% optimal length) over a wide range of bandwidths (4-90Hz). When active muscle is stretched appreciably, titin can develop enormous forces [7], [8], which may prevent further lengthening and injury. The stiffness that best captures the response of muscle to the small perturbations of Kirsch et al. [5] is far greater than the stiffness that best captures the response of muscle to large perturbations [7], [8]. Since everyday movements are often accompanied by both large and small kinematic perturbations, it is important to accurately capture these two processes.

However, there is likely no single muscle model that can replicate the force response of muscle to small [5] and large perturbations [7], [8] while also retaining the capability to reproduce the experiments of Hill [9] and Gordon et al. [10]. Unfortunately, this means that simulation studies that depend on an accurate representation of muscle impedance may reach conclusions well justified in simulation but not in reality. In this work, we focus on formulating a mechanistic muscle model2 that can replicate the force response of active muscle to length perturbations both great and small.

There are predominantly three classes of models that are used to simulate musculoskeletal responses: phenomenological models constructed using Hill’s famous force-velocity relationship [9], mechanistic Huxley [11]–[13] models in which individual elastic crossbridges are incorporated, and linearized muscle models [14], [15] which are accurate for small changes in muscle length. Kirsch et al. [5] demon-strated that, for small perturbations, the force response of muscle is well represented by a spring in parallel with a damper. Neither Hill nor Huxley models are likely to replicate Kirsch et al.’s [5] experiments because a Hill muscle model [16], [17] does not contain any active spring elements; while a Huxley model lacks an active damping element. Although linearized muscle models can replicate Kirsch et al.’s experiment [5], these models are only accurate for small changes in length and cannot replicate the Hill’s nonlinear force-velocity relation [9], nor Gordon et al.’s [10] nonlinear force-length relation. However, there have been significant improvements to the canonical forms of phenomenological, mechanistic, and linearized muscle models that warrant closer inspection.

Several novel muscle models have been proposed to improve upon the accuracy of Hill-type muscle models during large active stretches. Forcinito et al. [18] modeled the velocity dependence of muscle using a rheological element3 and an elastic rack rather than embedding the force-velocity relationship in equations directly, as is done in a typical Hill model [16], [17]. This modification allows Forcinito et al.’s [18] model to more faithfully replicate the force development of active muscle, as compared to a Hill-type model, during ramp length changes of ≈ 10%4 of the optimal CE length, and across velocities of 4 − 11% of the maximum contraction velocity5. Haeufle et al. [20] made use of a serial-parallel network of spring-dampers to allow their model to reproduce Hill’s force-velocity relationship [9] mechanistically rather than embedding the experimental curve directly in their model. This modification allowed Haeufle et al.’s model to simulate high speed reaching movements that agree more closely with experimental data [21] than is possible with a typical Hill model. Gü nther et al. [22] evaluated how accurately a variety of spring-damper models were able to reproduce the microscopic increases in crossbridge force in response to small length changes. While each of these models improves upon the force response of the Hill model to ramp length changes, none are likely to reproduce Kirsch et al.’s experiment [5] because the linearized versions of these models lead to a serial, rather than a parallel, connection of a spring and a damper: Kirsch et al. [5] specifically showed (see Figure 3 of [5]) that a serial connection of a spring-damper fails to reproduce the phase shift between force and length present in their experimental data.

Titin [23], [24] has been more recently investigated to explain how lengthened muscle can develop active force when lengthened both within, and beyond, actin-myosin overlap [8]. Titin is a gigantic multi-segmented protein that spans a half-sarcomere, attaching to the Z-line at one end and the middle of the thick filament at the other end [25]. In skeletal muscle, the two sections nearest to the Z-line, the proximal immunoglobulin (IgP) segment and the PEVK segment — rich in the amino acids proline (P), glutamate (E), valine (V) and lysine (K) — are the most compliant [26] since the distal immunoglobulin (IgD) segments bind strongly to the thick filament [27]. Titin has proven to be a complex filament, varying in composition and geometry between different muscle types [28], [29], widely between species [30], and can apply activation dependent forces to actin [31]. It has proven challenging to determine which interactions dominate between the various segments of titin and the other filaments in a sarcomere. Experimental observations have reported titin-actin interactions at myosin-actin binding sites [32], [33], between titin’s PEVK region and actin [34], [35], between titin’s N2A region and actin [36], and between the PEVK-IgD regions of titin and myosin [37]. This large variety of experimental observations has led to a correspondingly large number of proposed hypotheses and models, most of which involve titin interacting with actin [38]–[43], and more recently with myosin [44].

The addition of a titin element to a model will result in more accurate force production during large active length changes, but does not affect the stiffness and damping of muscle at modest sarcomere lengths because of titin’s relatively low stiffness. At sarcomere lengths of or less, the stiffness of the actin-myosin load path with a single attached crossbridge (0.22 − 1.15 pN/nm) equals or exceeds the stiffness of 6 passive titin filaments (0.0348 − 0.173 pN/nm), and our estimated stiffness of 6 active titin filaments (0.0696 − 0.346 pN/nm, see Appendix A for further details). When fully activated, the stiffness of the actin-myosin load path (4.05 − 18.4 pN/nm) far exceeds that of both the passive titin (0.0348 − 0.173 pN/nm), and our estimated active titin (0.0696 − 0.346 pN/nm) load paths. Since titin-focused models have not made any changes to the modeled myosin-actin interaction beyond a Hill [16], [17] or Huxley [11], [12] model, it is unlikely that these models would be able to replicate Kirsch et al.’s experiments [5].

Although most motor control simulations [2], [45]–[48] make use of the canonical linearized muscle model, phenomenological muscle models have also been used and modified to include stiffness. Sartori et al. [49] modeled muscle stiffness by evaluating the partial derivative of the force developed by a Hill-type muscle model with respect to the contractile element (CE) length. Although this approach is mathematically correct, the resulting stiffness is heavily influenced by the shape of the force-length curve and can lead to inaccurate results: at the optimal CE length this approach would predict an active muscle stiffness of zero since the slope of the force-length curve is zero; on the descending limb this approach would predict a negative active muscle stiffness since the slope of the force-length curve is negative. In contrast, CE stiffness is large and positive near the optimal length [5], and there is no evidence for negative stiffness on the descending limb of the force-length curve [7]. Although the stiffness of the CE can be kept positive by shifting the passive force-length curve, which is at times used in finite-element-models of muscle [50], this introduces a new problem: the resulting passive CE stiffness cannot be lowered to match a more flexible muscle. In contrast, De Groote et al. [51], [52] modeled short-range-stiffness using a stiff spring in parallel with the active force element of a Hill-type muscle model. While the approach of De Groote et al. [51], [52] likely does improve the response of a Hill-type muscle model for small perturbations, there are several drawbacks: the short-range-stiffness of the muscle sharply goes to zero outside of the specified range whereas in reality the stiffness is only reduced [5] (see Fig. 9A); the damping of the canonical Hill-model has been left unchanged and likely differs substantially from biological muscle [5].

In this work, we propose a model that can capture the force development of muscle to perturbations that vary in size and timescale, and yet is described using only a few states making it well suited for large-scale simulations. When active, the response of the model to perturbations within actin-myosin overlap is dominated by a viscoelastic crossbridge element that has different dynamics across time-scales: over brief time-scales the viscoelasticity of the lumped crossbridge dominates the response of the muscle [5], while over longer time-scales the force-velocity [9] and force-length [10] properties of muscle dominate. To capture the active forces developed by muscle beyond actin-myosin overlap we added an active titin element which, similar to existing models [38], [40], features an activation-dependent6 interaction between titin and actin. To ensure that the various parts of the model are bounded by reality, we have estimated the physical properties of the viscoelastic crossbridge element as well as the active titin element using data from the literature.

While our main focus is to develop a more accurate muscle model, we would like the model to be well suited to simulating systems that contain tens to hundreds of muscles. Although Huxley models have been used to simulate whole-body movements such as jumping [53], the memory and processing requirements associated with simulating a single muscle with thousands of states is high. Instead of modeling the force development of individual crossbridges, we lump all of the crossbridges in a muscle together so that we have a small number of states to simulate per muscle.

To evaluate the proposed model, we compare simulations of experiments to original data. We examine the response of active muscle to small perturbations over a wide band-width by simulating the stochastic perturbation experiments of Kirsch et al. [5]. Herzog et al.’s [7] active-lengthening experiments are used to evaluate the response of the model when it is actively lengthened within actin-myosin overlap. Next, we use Leonard et al.’s [8] active lengthening experiments to see how the model compares to reality when it is actively lengthened beyond actin-myosin overlap. In addition, we examine how well the model can reproduce the force-velocity experiments of Hill [9] and force-length experiments of Gordon et al. [10]. Since Hill-type models are so commonly used, we also replicate all of the simulated experiments using Millard et al.’s [17] Hill-type muscle model to make the differences between these two types of models clear.

Model

We begin by treating whole muscle as a scaled half-sarcomere that is pennated at an angle α with respect to a tendon (Fig. 1A). The assumption that mechanical properties scale with size is commonly used when modeling muscle [16] and makes it possible to model vastly different musculotendon units (MTUs) by simply changing the architectural and contraction properties: the maximum isometric force , the optimal CE length (at which the CE develops ), the pennation angle αo of the CE (at a length of ) with respect to the tendon, the maximum shortening velocity of the CE, and the slack length of the tendon . Many properties of sarcomeres scale with and scales with physiological cross-sectional area [54], the force-length property scales with [55], the maximum normalized shortening velocity of different CE types scales with across animals great and small [56], and titin’s passive-force-length properties scale from single molecules to myofibrils [57], [58]

The name of the VEXAT model comes from the viscoelastic crossbridge and active titin elements (A.) in the model. Active tension generated by the lumped crossbridge flows through actin, myosin, and the adjacent sarcomeres to the attached tendon (B.). Titin is modeled as two springs of length 𝓁 1 and 𝓁 2 in series with the rigid segments LT12 and LIgD. Viscous forces act between titin and actin in proportion to the activation of the muscle (C.), which reduces to negligible values in a purely passive muscle (D.). We modeled actin and myosin as rigid elements; the XE, titin, and the tendon as viscoelastic elements; and the ECM as an elastic element.

The proposed model has several additional properties that we assume scale with and inversely with the maximum active isometric stiffness and damping , the passive forces due to the extracellular matrix (ECM), and passive forces due to titin. As crossbridge stiffness is well studied [59], we assume that muscle stiffness due to crossbridges scales such that where is the maximum normalized stiffness. This scaling is just what would be expected when many crossbridges [59] act in parallel across the cross-sectional area of the muscle, and act in series along the length of the muscle. Although the intrinsic damping properties of crossbridges are not well studied, we assume that the linear increase in damping with activation observed by Kirsch et al. [5] is due to the intrinsic damping properties of individual cross-bridges which will also scale linearly with and inversely with where is the maximum normalized damping. For the remainder of the paper, we refer to the proposed model as the VEXAT model due to the viscoelastic (VE) crossbridge (X) and active-titin (AT) elements of the model.

To reduce the number of states needed to simulate the VEXAT model, we lump all of the attached crossbridges into a single lumped crossbridge element (XE) that attaches at 𝓁 S (Fig. 1A) and has intrinsic stiffness and damping properties that vary with the activation and force-length properties of muscle. The active force developed by the XE at the attachment point to actin is transmitted to the main myosin filament, the M-line, and ultimately to the tendon (Fig. 1B). In addition, since the stiffness of actin [60] and myosin filaments [61] greatly exceeds that of crossbridges [62], we treat actin and myosin filaments as rigid to reduce the number of states needed to simulate this model. Similarly, we have lumped the six titin filaments per half-sarcomere (Fig. 1A) together to further reduce the number of states needed to simulate this model.

The addition of a titin filament to the model introduces an additional active load-path (Fig. 1C) and an additional passive load-path (Fig. 1D). As is typical [16], [17], we assume that the passive elasticity of these structures scale linearly with and inversely with . Since the VEXAT model has two passive load paths (Fig. 1D), we further assume that the proportion of the passive force due to the extra-cellular-matrix (ECM) and titin does not follow a scale dependent pattern, but varies from muscle-to-muscle as observed by Prado et al. [58].

Earlier version: As previously mentioned, there are several theories to explain how titin interacts with the other filaments in activated muscle. While there is evidence for titin-actin interaction near titin’s N2A region [36], there is also support for a titin-actin interaction occurring near titin’s PEVK region [34], [35], and for a titin-myosin interaction near the PEVK-IgD region [37]. For the purposes of our model, we will assume a titin-actin interaction because current evidence weighs more heavily towards a titin-actin interaction than a titin-myosin interaction. Next, we assume that the titin-actin interaction takes place somewhere in the PEVK segment for two reasons: first, there is evidence for a titin-actin interaction [34], [35] in the PEVK segment; and second, there is evidence supporting an interaction at the proximal end of the PEVK segment (N2A-actin interaction) [36]. We have left the point within the PEVK segment that attaches to actin as a free variable since there is some uncertainty about what part of the PEVK segment interacts with actin.

Revised version: As previously mentioned, there are several theories to explain how titin interacts with the other filaments in activated muscle. While there is evidence for titin-actin interaction near titin’s N2A region [36], there is also support for a titin-actin interaction occurring near titin’s PEVK region [34], [35], and for a titin-myosin interaction near the PEVK-IgD region [37]. For the purposes of our model, we will assume a titin-actin interaction because current evidence weighs more heavily towards a titin-actin interaction than a titin-myosin interaction. Next, we assume that the titin-actin interaction takes place somewhere in the PEVK segment for two reasons: first, there is evidence for a titin-actin interaction [34], [35] in the PEVK segment; and second, there is evidence supporting an interaction at the proximal end of the PEVK segment (N2A-actin interaction) [36]. We have left the point within the PEVK segment that attaches to actin as a free variable since there is some uncertainty about what part of the PEVK segment interacts with actin. The nature of the mechanical interaction between titin and the other filaments in an active sarcomere remains uncertain. Here we assume that this interaction is not a rigid attachment, but instead is an activation dependent damping to be consistent with the observations of Kellermayer and Granzier [31] and Dutta et al. [36]: adding titin filaments and calcium slowed, but did not stop, the progression of actin filaments across a plate covered in active crossbridges (an in-vitro motility assay). When activated, we assume that the amount of damping between titin and actin scales linearly with and inversely with

Removed in the revised version: The nature of the mechanical interaction between titin and the other filaments in an active sarcomere remains uncertain. Here we assume that this interaction is not a rigid attachment, but instead is an activation dependent damping to be consistent with the observations of Kellermayer and Granzier [31] and Dutta et al. [36]: adding titin filaments and calcium slowed, but did not stop, the progression of actin filaments across a plate covered in active crossbridges (an in-vitro motility assay). When activated, we assume that the amount of damping between titin and actin scales linearly with and inversely with .

After lumping all of the crossbridges and titin filaments together we are left with a rigid-tendon MTUmodel that has two generalized positions and an elastic-tendon MTUmodel that has three generalized positions Given these generalized positions, the path length 𝓁 P, and a pennation model, all other lengths in the model can be calculated. Here we use a constant thickness pennation model to evaluate the pennation angle of a CE with a rigid-tendon, and to evaluate the pennation angle of a CE with an elastic-tendon. We have added a small compressive element KE (Fig. 1A) to prevent the model from reaching the numerical singularity that exists as approaches , the length at which α → 90° in Eqns. 6 and 7. The tendon length of an elastic-tendon model is the difference between the path length and the CE length along the tendon. The length of the XE is the difference between the half-sarcomere length and the sum of the average point of attachment 𝓁 S and the length of the myosin filament LM. The length of 𝓁 2, the lumped PEVK-IgD segment, is the difference between the half-sarcomere length and the sum of the length from the Z-line to the actin binding site on titin (𝓁 1) and the length of the IgD segment that is bound to myosin (LIgD). Finally, the length of the extra-cellular-matrix 𝓁 ECM is simply half the length of the CE since we are modeling a half-sarcomere.

We have some freedom to choose the state vector of the model and the differential equations that define how the muscle responds to length and activation changes. The experiments we hope to replicate depend on phenomena that take place at different time-scales: Kirsch et al.’s [5] stochastic perturbations evolve over brief time-scales, while all of the other experiments take place at much longer time-scales. Here we mathematically decouple phenomena that affect brief and long time-scales by making a second-order model that has states of the average point of crossbridge attachment 𝓁 S, and velocity v S. When the activation a state and the titin-actin interaction model are included, the resulting rigid-tendon model that has a total of four states and the elastic-tendon model has five states. For the purpose of comparison, a Hill-type muscle model with a rigid-tendon has a single state (a), while an elastic-tendon model has two states (a and 𝓁 M) [17].

Earlier version: Before proceeding, a small note on notation: through-out this work we will use an underbar to indicate a vector, bold font to indicate a curve, a tilde for a normalized quantity, and a capital letter to indicate a constant. Unless indicated otherwise, curves are constructed using 𝒞2 continuous7 Bézier splines so that the model is compatible with gradient-based optimization. Normalized quantities within the CE follow a specific convention: lengths and velocities are normalized by the optimal CE length , forces by the maximum active isometric tension , stiffness and damping by . Velocities used as input to the force-velocity relation f V are further normalized by and annotated using a hat: . Tendon lengths and velocities are normalized by tendon slack length, while forces are normalized by .

Revised version: Before proceeding, a small note on notation: throughout this work we will use an underbar to indicate a vector, bold font to indicate a curve, a tilde for a normalized quantity, and a capital letter to indicate a constant. Unless indicated otherwise, curves are constructed using C2 continuous7 Bézier splines so that the model is compatible with gradient-based optimization. Normalized quantities within the CE follow a specific convention: lengths and velocities are normalized by the optimal CE length , forces by the maximum active isometric tension , stiffness and damping by . Velocities used as input to the force-velocity relation f V are further normalized by annotated using a hat: . Tendon lengths and velocities are normalized by tendon slack length, while forces are normalized by .

Earlier version: To evaluate the state derivative of the model, we require equations for , and v M if the tendon is elastic. For we use of the first order activation dynamics model described in Millard et al. [17]8 which uses a lumped first order ordinary-differential-equation (ODE) to describe how a fused tetanus electrical excitation leads to force development in an isometric muscle. We formulated the equation for with the intention of having the model behave like a spring-damper over small time-scales, but to converge to the tension developed by a Hill-type model over longer time-scales, where f L(·) is the active-force-length curve (Fig. 2A), f PE(·) is the passive-force-length curve (Fig. 2A), and f V(·) is the force-velocity (Fig. 2B). The normalized tension developed by the VEXAT model differs from that of a Hill model, Eqn. 14, because it has no explicit dependency on , includes four passive terms, and a lumped viscoelastic crossbridge element. The four passive terms come from the ECM element (Fig. 3A), the PEVK-IgD element (Fig. 3A and B), the compressive term (prevents cos α from reaching a length of 0), and a numerical damping term (where is small). The active force developed by the XE’s lumped crossbridge is scaled by the fraction of the XE that is active and attached, , where f L(·) is the active-force-length relation (Fig. 2A). We evaluate f L using in Eqn. 15, rather than as in Eqn. 14, since actin-myosin overlap is independent of crossbridge strain. With derived, we can proceed to model the acceleration of the CE, , so that it is driven over time by the force imbalance between the XE’s active tension and that of a Hill model.

Revised version: To evaluate the state derivative of the model, we require equations for and v 1, and v M if the tendon is elastic. For we use of the first order activation dynamics model described in Millard et al. [17]8 which uses a lumped first order ordinary-differential-equation (ODE) to describe how a fused tetanus electrical excitation leads to force develop-ment in an isometric muscle. We formulated the equation for with the intention of having the model behave like a spring-damper over small time-scales, but to converge to the tension developed by a Hill-type model over longer time-scales, where f L(·) is the active-force-length curve (Fig. 2A), f PE(·) is the passive-force-length curve (Fig. 2A), and f V(·) is the force-velocity (Fig. 2B).

The model relies on Bézier curves to model the nonlinear effects of the active-force-length curve, the passive-force-length curves (A.), and the force-velocity curve (B.). Since nearly all of the reference experiments used in Sec. 3 have used cat soleus, we have fit the active-force-length curve (f L(·)) and passive-force-length curves (f PE(·)) to the cat soleus data of Herzog and Leonard 2002 [7]. The concentric side of the force-velocity curve (f V(·)) has been fitted to the cat soleus data of Herzog and Leonard 1997 [63].

The passive force-length curve has been decomposed such that 56% of it comes from the ECM while 44% comes from titin to match the average of ECM-titin passive force distribution (which ranges from 43%-76%) reported by Prado et al. [58] (A.). The elasticity of the titin segment has been further decomposed into two serially connected sections: the proximal section consisting of the T12, proximal IgP segment and part of the PEVK segment, and the distal section consisting of the remaining PEVK section and the distal Ig segment (B.). The stiffness of the IgP and PEVK segments has been chosen so that the model can reproduce the movements of IgP/PEVK and PEVK/IgD boundaries that Trombitás et al. [26] (C.) observed in their experiments. The curves that appear in subplots A. and B. come from scaling the two-segmented human soleus titin model to cat soleus muscle. The curves that appear in subplot C compare the human soleus titin model’s IgP, PEVK, and IgD force-length relations to the data of Trombitás et al. [26] (see in Appendix B for details).

Added in the revised version: The normalized tension developed by the VEXAT model differs from that of a Hill model, Eqn. 14, because it has no explicit dependency on , includes four passive terms, and a lumped viscoelastic crossbridge element. The four passive terms come from the ECM element (Fig. 3A), the PEVK-IgD element (Fig. 3A and B), the compressive term (prevents cos α from reaching a length of 0), and a numerical damping term (where is small). The active force developed by the XE’s lumped crossbridge is scaled by the fraction of the XE that is active and attached, , where f L(·) is the active-force-length relation (Fig. 2A). We evaluate f L using in Eqn. 15, rather than as in Eqn. 14, since actin-myosin overlap is independent of crossbridge strain. With derived, we can proceed to model the acceleration of the CE, , so that it is driven over time by the force imbalance between the XE’s active tension and that of a Hill model.

We set the first term of so that, over time, the CE is driven to develop the same active tension as a Hill-type model [17] (terms highlighted in blue) where τ S is a time constant and is the force-velocity curve (Fig. 2B). The rate of adaptation of the model’s tension, to the embedded Hill model, is set by the time constant τ S: as τ S is decreased the VEXAT model converges more rapidly to a Hill-type model; as τ S is increased the active force produced by the model will look more like a spring-damper. Our preliminary simulations indicate that there is a trade-off to choosing τ S: when τ S is large the model will not shorten rapidly enough to replicate Hill’s experiments, while if τ S is small the low-frequency response of the model is compromised when Kirsch et al.’s [5] experiments are simulated.

Earlier version: The remaining two terms, and , have been included for numerical reasons specific to this model formulation rather than muscle physiology. We include a term that damps the rate of actinmyosin translation, , to prevent this second-order system from unrealistically oscillating9. The final term , where GL and GV are scalar gains, and aL is a low-activation threshold (aL is 0.05 in this work). This final term has been included as a consequence of the generalized positions we have chosen. When the CE is nearly deactivated (as a approaches aL), this term forces and to shadow the location and velocity of the XE attachment point. This ensures that if the XE is suddenly activated, that it attaches with little strain. We had to include this term because we made ℓ S a state variable, rather than ℓ X. We chose ℓ S as a state variable, rather than ℓ X, so that the states are more equally scaled for numerical integration.

Revised version: The remaining two terms, and , have been included for numerical reasons specific to this model formulation rather than muscle physiology. We include a term that damps the rate of actin-myosin translation, , to prevent this second-order system from unrealistically oscillating9. The final term , where GL and GV are scalar gains, and aL is a low-activation threshold (aL is 0.05 in this work). This final term has been included as a consequence of the generalized positions we have chosen. When the CE is nearly deactivated (as a approaches aL), this term forces and to shadow the location and velocity of the XE attachment point. This ensures that if the XE isbsuddenly activated, that it attaches with little strain. We had to include this term because we made 𝓁 S a state variable, rather than 𝓁 X. We chose 𝓁 S as a state variable, rather than 𝓁 X, so that the states are more equally scaled for numerical integration.

Earlier version: The passive force developed by the CE in Eqn. 15 is the sum of the elastic forces (Fig. 3A) developed by the force-length curves of titin and and the ECM . We model titin’s elasticity as being due to two serially connected elastic segments: the first elastic segment is formed by lumping together the IgP segment and a fraction Q of the PEVK segment, while the second elastic segment is formed by lumping together the remaining (1 − Q) of the PEVK segment with the free IgD section. Our preliminary simulations of Herzog and Leonard’s active lengthening experiment [7] indicate that a Q value of 0.5, positioning the PEVK-actin attachment point that is near the middle of the PEVK segment, allows the model to develop sufficient tension when actively lengthened. The large section of the IgD segment that is bound to myosin is treated as rigid.

Revised version: The passive force developed by the CE in Eqn. 15 is the sum of the elastic forces (Fig. 3A) developed by the force-length curves of titin and and the . We model titin’s elasticity as being due to two serially connected elastic segments: the first elastic segment is formed by lumping together the IgP segment and a fraction Q of the PEVK segment, while the second elastic segment is formed by lumping together the remaining (1 − Q) of the PEVK segment with the free IgD section. Our preliminary simulations of Herzog and Leonard’s active lengthening experiment [7] indicate that a Q value of 0.5, positioning the PEVK-actin attachment point that is near the middle of the PEVK segment, allows the model to develop sufficient tension when actively lengthened. The large section of the IgD segment that is bound to myosin is treated as rigid.

The curves that form , and have been carefully constructed to satisfy three experimental observations: that the total passive force-length curve of titin and the ECM match the observed passive force-length curve of the muscle (Fig. 2A and Fig. 3A) [58]; that the proportion of the passive force developed by titin and the ECM is within experimental observations [58] (Fig. 3A); and that the Ig domains and PEVK residues show the same relative elongation as observed by Trombitás et al. [26] (Fig. 3C). Even though Trombitás et al.’s [26] measurements come from human soleus titin, we can construct the force-length curves of other titin isoforms if the number of proximal Ig domains, PEVK residues, and distal Ig domains are known (see Appendix B.3). Since the passive–force-length relation and the results of Trombitás et al. [26] are at modest lengths, we consider two different extensions to the force-length relation to simulate extreme lengths: first, a simple linear extrapolation; second, we extend the force-length relation of each of titin’s segments to follow a worm-like-chain (WLC) model [65] (see Appendix B.3 for details on the WLC model). With titin’s passive force-length relations defined, we can next consider what happens to titin in active muscle.

Removed in the revised version: When active muscle is lengthened, it produces an enhanced force that persists long after the lengthening has ceased called residual force enhancement (RFE) [7]. For the purposes of the VEXAT model, we assume that RFE is produced by titin. Experiments have demonstrated RFE on both the ascending limb [66] and descending limb of the force-length [7] relation. The amount of RFE depends both on the final length of the stretch [67] and the magnitude of the stretch: on the ascending limb the amount of RFE varies with the final length but not with stretch magnitude, while on the descending limb RFE varies with stretch magnitude.

Earlier version: To develop RFE, we assume that some point of the PEVK segment bonds with actin through an activation-dependent damper. The VEXAT model’s distal segment of titin, ℓ 2, can contribute to RFE when the titin-actin bond is formed and CE is lengthened beyond , the shortest CE length at which the PEVK-actin bond can form. In this work, we set to be equal to the slack length of the CE’s force-length relation (see Table 1E and H). To incorporate the asymmetric length dependence of RFE [67], we introduce a smooth step-up function that transitions from zero to one as extends beyond , where the sharpness of the transition is controlled by R. Similar to Hisey et al.’s experimental work [67], active lengthening on the ascending limb will produce similar amounts of RFE since and the titin-actin bond is prevented from forming below . In contrast, the amount of RFE on the descending limb increases with increasing stretch magnitudes since the titin-actin bond is able to form across the entire descending limb.

Revised version: When active muscle is lengthened, it produces an enhanced force that persists long after the lengthening has ceased called residual force enhancement (RFE) [7]. For the purposes of the VEXAT model, we assume that RFE is produced by titin. Experiments have demonstrated RFE on both the ascending limb [66] and descending limb of the force-length [7] relation. The amount of RFE depends both on the final length of the stretch [67] and the magnitude of the stretch: on the ascending limb the amount of RFE varies with the final length but not with stretch magnitude, while on the descending limb RFE varies with stretch magnitude. To develop RFE, we assume that some point of the PEVK segment bonds with actin through an activation-dependent damper. The VEXAT model’s distal segment of titin, 𝓁 2, can contribute to RFE when the titin-actin bond is formed and CE is lengthened beyond , the shortest CE length at which the PEVK-actin bond can form. In this work, we set to be equal to the slack length of the CE’s force-length relation (see Table 1E and H). To incorporate the asymmetric length dependence of RFE [67], we introduce a smooth step-up function that transitions from zero to one as extends beyond , where the sharpness of the transition is controlled by R. Similar to Hisey et al.’s experimental work [67], active lengthening on the ascending limb will produce similar amounts of RFE since and the titin-actin bond is prevented from forming below . In contrast, the amount of RFE on the descending limb increases with increasing stretch magnitudes since the titin-actin bond is able to form across the entire descending limb.

At very long CE lengths, the modeled titin-actin bond can literally slip off of the end of the actin filament (Fig. 1A) when the distance between the Z-line and the bond, , exceeds the length of the actin filament, . To break the titin-actin bond at long CE lengths we introduce a smooth step-down function The step-down function uL transitions from one to zero when the titin-actin bond reaches , the end of the actin filament.

The strength of the titin-actin bond also appears to vary nonlinearly with activation. Fukutani and Herzog [68] observed that the absolute RFE magnitude produced by actively lengthened fibers is similar between normal and reduced contractile force states. Since these experiments [68] were performed beyond the optimal CE length, titin could be contributing to the observed RFE as previously described. The consistent pattern of absolute RFE values observed by Fukutani and Herzog [68] could be produced if the titin-actin bond saturated at its maximum strength even at a reduced contractile force state. To saturate the titin-actin bond, we use a final smooth step function where A° is the threshold activation level at which the bond saturates. While we model the strength of the titin-actin bond as being a function of activation, which is proportional Ca2+ concentration [69], this is a mathematical convenience. The work of Leonard et al. [8] makes it clear that both Ca2+ and crossbridge cycling are needed to allow titin to develop enhanced forces during active lengthening: no enhanced forces are observed in the presence of Ca2+ when crossbridge cycling is chemically inhibited. Putting this all together, the active damping acting between the titin and actin filaments is given by the product of where is the maximum damping coefficient.

Earlier version: With a model of the titin-actin bond derived, we can fo-cus on how the bond location moves in response to applied forces. Since we are ignoring the mass of the titin filament, the PEVK-attachment point is balanced by the forces applied to it and the viscous forces developed between titin and actin due to the active and a small amount of passive damping . Since Eqn. 20 is linear in , we can solve directly for it The assumption of whether the tendon is rigid or elastic affects how the state derivative is evaluated and how expensive it is to compute. While all of the position dependent quantities can be evaluated using Eqns. 6-11 and the generalized positions, evaluating the generalized velocities of a rigid-tendon and elastic-tendon model differ substantially. The CE velocity v M and pennation angular velocity of a rigid-tendon model can be evaluated directly given the path length, velocity, and the time derivatives of Eqns. 6 and 8. After v 1 is evaluated using Eqn. 21, the velocities of the remaining segments can be evaluated using the time derivatives of Eqns. 9-11.

Revised version: With a model of the titin-actin bond derived, we can focus on how the bond location moves in response to applied forces. Since we are ignoring the mass of the titin filament, the PEVK-attachment point is balanced by the forces applied to it and the viscous forces developed between titin and actin due to the active and a small amount of passive damping Since Eqn. 20 is linear in , we can solve directly for it The assumption of whether the tendon is rigid or elastic affects how the state derivative is evaluated and how expensive it is to compute. While all of the position dependent quantities can be evaluated using Eqns. 6-11 and the generalized positions, evaluating the generalized velocities of a rigid-tendon and elastic-tendon model differ substantially.

Added in the revised version: The CE velocity v M and pennation angular velocity of a rigid-tendon model can be evaluated directly given the path length, velocity, and the time derivatives of Eqns. 6 and 8. After v 1 is evaluated using Eqn. 21, the velocities of the remaining segments can be evaluated using the time derivatives of Eqns. 9-11.

Earlier version: Evaluating the CE rate of lengthening, v M, for an elastictendon muscle model is more involved. As is typical of lumped parameter muscle models [16], [17], [70], here we assume that difference in tension, , between the CE and the tendon is negligible10. During our preliminary simulations it became clear that treating the tendon as an idealized spring degraded the ability of the model to replicate the experiment of Kirsch et al. [5] particularly at high frequencies. Kirsch et al. [5] observed a linear increase in the gain and phase profile between the output force and the input perturbation applied to the muscle. This pattern in gain and phase shift can be accurately reproduced by a spring in parallel with a damper. Due to the way that impedance combines in series11, the models of both the CE and the tendon need to have parallel spring and damper elements so that the entire MTU, when linearized, appears to be a spring in parallel with a damping element. We model tendon force using a nonlinear spring and damper model where the damping coefficient , is a linear scal-ing of the normalized tendon stiffness by U, a constant scaling coefficient. We have chosen this specific damping model because it fits the data of Netti et al. [72] and captures the structural coupling between tendon stiffness and damping (see Appendix B.1 and Fig. 15 for further details).

Revised version: Evaluating the CE rate of lengthening, v M, for an elastictendon muscle model is more involved. As is typical of lumped parameter muscle models [16], [17], [70], here we assume that difference in tension, , between the CE and the tendon is negligible10. During our preliminary simulations it became clear that treating the tendon as an idealized spring degraded the ability of the model to replicate the experiment of Kirsch et al. [5] particularly at high frequencies. Kirsch et al. [5] observed a linear increase in the gain and phase profile between the output force and the input perturbation applied to the muscle. This pattern in gain and phase shift can be accurately reproduced by a spring in parallel with a damper. Due to the way that impedance combines in series 11, the models of both the CE and the tendon need to have parallel spring and damper elements so that the entire MTU, when linearized, appears to be a spring in parallel with a damping element. We model tendon force using a nonlinear spring and damper model where the damping coefficient , is a linear scal, ing of the normalized tendon stiffness by U, a constant scaling coefficient. We have chosen this specific damping model because it fits the data of Netti et al. [72] and captures the structural coupling between tendon stiffness and damping (see Appendix B.1 and Fig. 15 for further details). Now that all of the terms in Eqn. 22 have been explicitly defined, we can use Eqn. 22 to solve for v M. Equation 22 becomes linear in v M after substituting the force models described in Eqns. 23 and 15, and the kinematic model equations described in Eqns. 8, 9 and 11 (along with the time derivatives of Eqns. 8-11). After some simplification we arrive at allowing us to evaluate the final state derivative in . During simulation the denominator of will always be finite since , and α < 90 ° due to the compressive element. The evaluation of in the VEXAT model is free of numerical singularities, giving it an advantage over a conventional Hill-type muscle model [17]. In addition, the VEXAT’s does not require iteration to numerically solve a root, giving it an advantage over a singularity-free formulation of the Hill model [17]. As with previous models, initializing the model’s state is not trivial and required the derivation of a model-specific method (see Appendix C for details).

Removed in the revised version: Now that all of the terms in Eqn. 22 have been explicitly defined, we can use Eqn. 22 to solve for v M. Equation 22 becomes linear in v M after substituting the force models described in Eqns. 23 and 15, and the kinematic model equations described in Eqns. 8, 9 and 11 (along with the time derivatives of Eqns. 8-11). After some simplification we arrive at allowing us to evaluate the final state derivative in . During simulation the denominator of will always be finite since , and α < 90° due to the compressive element. The evaluation of in the VEXAT model is free of numerical singularities, giving it an advantage over a conventional Hill-type muscle model [17]. In addition, the VEXAT’s does not require iteration to numerically solve a root, giving it an advantage over a singularity-free formulation of the Hill model [17]. As with previous models, initializing the model’s state is not trivial and required the derivation of a model-specific method (see Appendix C for details).

Biological Benchmark Simulations

Earlier version: In order to evaluate the model, we have selected three experiments that capture the responses of active muscle to small, medium, and large length changes. The small (1-3.8% ) stochastic perturbation experiment of Kirsch et al. [5] demonstrates that the impedance of muscle is well described by a stiff spring in parallel with a damper, and that the spring-damper coefficients vary linearly with active force. The active impedance of muscle is such a fundamental part of motor learning that the amount of impedance, as indicated by co-contraction, is used to define how much learning has actually taken place [1], [77]: co-contraction is high during initial learning, but decreases over time as a task becomes familiar. The active lengthening experiment of Herzog and Leonard [7] shows that modestly stretched (7-21% ) biological muscle has positive stiffness even on the descending limb of the active force-length curve . In contrast, a conventional Hill model [16], [17] can have negative stiffness on the descending limb of the active-force-length curve, a property that is both mechanically unstable and unrealistic. The final active lengthening experiment of Leonard et al. [8] unequivo-cally demonstrates that the CE continues to develop active forces during extreme lengthening (329% ) which exceeds actin-myosin overlap. Active force development beyond actin-myosin overlap is made possible by titin, and its activation dependent interaction with actin [8]. The biological benchmark simulations conclude with a replication of the force-velocity experiments of Hill [9] and the force-length experiments of Gordon et al. [10].

Revised version: In order to evaluate the model, we have selected three experiments that capture the responses of active muscle to small, medium, and large length changes. The small stochastic perturbation experiment of Kirsch et al. [5] demonstrates that the impedance of muscle is well described by a stiff spring in parallel with a damper, and that the spring-damper coefficients vary linearly with active force. The active impedance of muscle is such a fundamental part of motor learning that the amount of impedance, as indicated by co-contraction, is used to define how much learning has actually taken place [1], [77]: co-contraction is high during initial learning, but decreases over time as a task becomes familiar. The active lengthening experiment of Herzog and Leonard [7] shows that modestly stretched biological muscle has positive stiffness even on the descending limb of the active force-length curve . In contrast, a conventional Hill model [16], [17] can have negative stiffness on the descending limb of the active-force-length curve, a property that is both mechanically unstable and unrealistic. The final active lengthening experiment of Leonard et al. [8] unequivocally demonstrates that the CE continues to develop active forces during extreme lengthening which exceeds actin-myosin overlap. Active force development beyond actin-myosin overlap is made possible by titin, and its activation dependent interaction with actin [8]. The biological benchmark simulations conclude with a replication of the force-velocity experiments of Hill [9] and the force-length experiments of Gordon et al. [10].

Earlier version: The VEXAT model requires the architectural muscle parameters needed by a conventional Hill-type muscle model as well as additional parameters. The additional parameters are needed for these component models: the lumped viscoelastic XE (Eqn. 1 and 2), XE-actin dynamics (Eqn. 16), the two-segment active titin model (Fig. 3), titin-actin dynamics (Eqns. 21), and the tendon damping model (Eqn. 23). Fortunately, there is enough experimental data in the literature that values can be found, fitted, or estimated directly for our simulations of experiments on cat soleus (Table 1), and rabbit psoas fibrils (see Appendix B for fitting details and Appendix H for rabbit psoas fibril model parameters).

Revised version: The VEXAT model requires the architectural muscle parameters (, and ) needed by a conventional Hill-type muscle model as well as additional parameters. The additional parameters are needed for these component models: the compressive element (Eqn. 15 and 24), the lumped viscoelastic XE (Eqn. 1 and 2), XE-actin dynamics (Eqn. 16), the two-segment active titin model (Fig. 3), titin-actin dynamics (Eqns. 21), and the tendon damping model (Eqn. 23). Fortunately, there is enough experimental data in the literature that values can be found, fitted, or estimated directly for our simulations of experiments on cat soleus (Table 1), and rabbit psoas fibrils (see Appendix B for fitting details and Appendix H for rabbit psoas fibril model parameters). The parameter values we have established for the cat soleus (Table 1 F.-I.) can serve as initial values when modeling other mammalian MTU’s because these parameters have been normalized (by , and where appropriate) and will scale appropriately given the architectural properties of a different MTU. By making use of these default values, the VEXAT model can be made to represent another MTU using exactly the same number of parameters as a Hill-type muscle model (Table 1A.-E.).

Stochastic Length Perturbation Experiments

Earlier version: In Kirsch et al.’s [5] in-situ experiment, the force response of a cat’s soleus muscle under constant stimulation was measured as its length was changed by small amounts. Kirsch et al. [5] applied stochastic length perturbations (Fig. 4A) to elicit force responses from the muscle (in this case a spring-damper Fig. 4B) across a broad range of frequencies (4-90 Hz) and across a range of small length perturbations (1-3.8% ). The resulting time-domain signals can be quite complicated (Fig. 4A and B) but contain rich measurements of how muscle transforms changes in length into changes in force.

Revised version: In Kirsch et al.’s [5] in-situ experiment, the force response of a cat’s soleus muscle under constant stimulation was measured as its length was changed by small amounts. Kirsch et al. [5] applied stochastic length perturbations (Fig. 4A) to elicit force responses from the muscle (in this case a spring-damper Fig. 4B) across a broad range of frequencies (4-90 Hz) and across a range of small length perturbations . The resulting time-domain signals can be quite complicated (Fig. 4A and B) but contain rich measurements of how muscle transforms changes in length into changes in force.

Evaluating a system’s gain and phase response begins by applying a pseudo-random input signal to the system and measuring its output (A). Both the input and output signals (A) are transformed into the frequency domain by expressing these signals as an equivalent sum of scaled and shifted sinusoids (simple example shown in B and C). Each individual input sinusoid is compared with the output sinusoid of the same frequency to evaluate how the system scales and shifts the input to the output (D). This process is repeated across all sinusoid pairs to produce a function that describes how an input sinusoid is scaled (E) and shifted (F) to an output sinusoid using only the measured data (A).

Added in the revised version: As long as muscle can be considered to be linear (a sinusoidal change in length produces a sinusoidal change in force), then system identification methods [78], [79] can be applied to extract a relationship between length x(t) and force y(t). We will give a brief overview of system identification methods here to make methods and results clearer. First, the time-domain signals (x(t) and y(t)) are transformed into an equivalent representation in the frequency-domain (X(s) and Y (s)) as a sum of scaled and shifted sine curves (Fig. 4B and C) using a Fourier transform [78].

Added in the revised version: In the frequency domain, we identify an LTI system of best fit H(s) that describes how muscle transforms changes in length into changes in force such that Y (s) = H(s) X(s). Next, we evaluate how H(s) scales the magnitude (gain) and shifts the timing (phase) of a sinusoid in X(s) into a sinusoid of the same frequency in Y (s) (Fig. 4D). This process is repeated across all frequency-matched pairs of input and output sinusoids to build a function of how muscle scales (Fig. 4E) and shifts (Fig. 4F) input length sinusoids into output force sinusoids. The resulting transformation turns two complicated time-domain signals (Fig. 4A) into a clear relationship in the frequency-domain that describes how muscle transforms length changes into force changes: a very slow (0Hz) length change will result in an output force that is scaled by 4.5 and is in phase (Fig. 4E and F), a 35 Hz sinusoidal length change will produce an output force that is scaled by 4.9 and leads the input signal by 24° (Fig. 4E and F), and frequencies between 0 Hz and 35 Hz will be between these two signals in terms of scaling and phase. These patterns of gain and phase can be used to identify a network of spring-dampers that is equivalent to the underlying linear system (the system in Fig. 4 A, E, and F is a 4.46 N/mm spring in parallel with a 0.0089 Ns/mm damper). Since experimental measurements often contain noise and small nonlinearities, the mathematical procedure used to estimate H(s) and the corresponding gain and phase profiles is more elaborate than we have described (see Appendix D for details).

Added in the revised version: Kirsch et al. [5] used system identification methods to identify LTI mechanical systems that best describes how muscle transforms input length waveforms to output force waveforms. The resulting LTI system, however, is only accurate when the relationship between input and output is approximately linear. Kirsch et al. [5] used the coherence squared, (Cxy)2, between the input and output to evaluate the degree of linearity: Y(s) is a linear transformation of X(s) at frequencies in which (Cxy)2 is near one, while Y (s) cannot be described as a linear function of X(s) at frequencies in which (Cxy)2 approaches zero. By calculating (Cxy)2 between the length perturbation and force waveforms, Kirsch et al. [5] identified the bandwidth in which the muscle’s response is approximately linear. Kirsch et al. [5] set the lower frequency of this band to 4 Hz, and Fig. 3 of Kirsch et al. [5] suggests that this corresponds to (Cxy)2 ≥ 0.67 though the threshold for (Cxy)2 is not reported. The upper frequency of this band was set to the cutoff frequency of the low-pass filter applied to the input (15, 35, or 90 Hz). Within this bandwidth, Kirsch et al. [5] compared the response of the specimen to several candidate models and found that a parallel spring-damper fit the muscle’s frequency response best. Next, they evaluated the stiffness and damping coefficients that best fit the muscle’s frequency response [5]. Finally, Kirsch et al. evaluated how much of the muscle’s time-domain response was captured by the spring-damper of best fit by evaluating the variance-accounted-for (VAF) between the two time-domain signals Astonishingly, Kirsch et al. [5] found that a spring-damper of best fit has a VAF of between 78-99%12 when compared to the experimentally measured forces f EXP. By repeating this experiment over a variety of stimulation levels (using both electrical stimulation and the crossed-extension reflex) Kirsch et al. [5] showed that these stiffness and damping coefficients vary linearly with the active force developed by the muscle. Further, Kirsch et al. [5] repeated the experiment using perturbations that had a variety of lengths (0.4 mm, 0.8mm, and 1.6mm) and bandwidths (15Hz, 35Hz, and 90Hz) and observed a peculiar quality of muscle: the damping coefficient of best fit increases as the bandwidth of the perturbation decreases (See Figures 3 and 10 of Kirsch et al. [5] for details). Here we simulate Kirsch et al.’s experiment [5] to determine, first, the VAF of the VEXAT model and the Hill model in comparison to a spring-damper of best fit; second, to compare the gain and phase response of the models to biological muscle; and finally, to see if the spring-damper coefficients of best fit for both models increase with active force in a manner that is similar to the cat soleus that Kirsch et al. studied [5].

Added in the revised version: To simulate the experiments of Kirsch et al. [5] we begin by creating the 9 stochastic perturbation waveforms used in the experiment that vary in perturbation amplitude (0.4mm, 0.8mm, and 1.6mm) and bandwidth (0-15 Hz, 0-35 Hz, and 0-90 Hz)13. The waveform is created using a vector that is composed of random numbers with a range of [−1, 1] that begins and ends with a series of zero-valued padding points.

Added in the revised version: Next, a forward pass of a 2nd order Butterworth filter is applied to the waveform and finally the signal is scaled to the appropriate amplitude (Fig. 5). The muscle model is then activated until it develops a constant tension at a length of . The musculotendon unit is then simulated as the length is varied using the previously constructed waveforms while activation is held constant. To see how impedance varies with active force, we repeated these simulations at ten evenly spaced tensions from 2.5N to 11.5N. Ninety simulations are required to evaluate the nine different perturbation waveforms at each of the ten tension levels. The time-domain length perturbations and force responses of the modeled muscles are used to evaluate the coherence squared of the signal, gain response, and phase responses of the models in the frequency-domain. Since the response of the models might be more nonlinear than biological muscle, we select a bandwidth that meets (Cxy)2 > 0.67 but otherwise follows the bandwidths analyzed by Kirsch et al. [5] (see Appendix D for details).

The perturbation waveforms are constructed by generating a series of pseudo-random numbers, padding the ends with zeros, by filtering the signal using a 2nd order low-pass filter (wave forms with -3dB cut-off frequencies of 90 Hz, 35 Hz and 15 Hz appear in A.) and finally by scaling the range to the desired limit (1.6mm in A.). Although the power spectrum of the resulting signals is highly variable, the filter ensures that the frequencies beyond the -3dB point have less than half their original power (B.).

Added in the revised version: When coupled with an elastic-tendon, the 15 Hz perturbations show that neither model can match the VAF of Kirsch et al.’s analysis [5] (compare Fig. 6A to G), while at 90Hz the VEXAT model reaches a VAF of 89% (Fig. 6D) which is within the range of 78-99% reported by Kirsch et al. [5]. In contrast, the Hill model’s VAF at 90 Hz remains low at 58% (Fig. 6J). While the VEXAT model has a gain profile in the frequency-domain that closer to Kirsch et al.’s data [5] than the Hill model (compare Fig. 6B to H and E to K), both models have a greater phase shift than Kirsch et al.’s data [5] at low frequencies (compare Fig. 6C to I and F to L). The phase response of the VEXAT model to the 90 Hz perturbation (Fig. 6F) shows the consequences of Eqn. 16: at low frequencies the phase response of the VEXAT model is similar to that of the Hill model, while at higher frequencies the model’s response becomes similar to a spring-damper. This frequency dependent response is a consequence of the first term in Eqn. 16: the value of τ S causes the response of the model to be similar to a Hill model at lower frequencies and mimic a spring-damper at higher frequencies. Both models show the same perturbation-dependent phase-response, as the damping coefficient of best fit increases as the perturbation bandwidth decreases: compare the damping coefficient of best fit for the 15Hz and 90Hz profiles for the VEXAT model (listed on Fig. 6A. and D.) and the Hill model (listed on Fig. 6G. and J., respectively).

The 15 Hz perturbations show that the VEXAT model’s performance is mixed: in the time-domain (A.) the VAF is lower than the 78-99% analyzed by Kirsch et al. [5]; the gain response (B.) follows the profile in Figure 3 of Kirsch et al. [5], while the phase response differs (C.). The response of the VEXAT model to the 90 Hz perturbations is much better: a VAF of 91% is reached in the time-domain (D.), the gain response follows the response of the cat soleus analyzed by Kirsch et al. [5], while the phase-response follows biological muscle closely for frequencies higher than 30 Hz. Although the Hill’s time-domain response to the 15 Hz signal has a higher VAF than the VEXAT model (G.), the RMSE of the Hill model’s gain response (H.) and phase response (I.) shows it to be more in error than the VEXAT model. While the VEXAT model’s response improved in response to the 90 Hz perturbation, the Hill model’s response does not: the VAF of the time-domain response remains low (J.), neither the gain (K.) nor phase responses (L.) follow the data of Kirsch et al. [5]. Note that the Hill model’s 90 Hz response was so nonlinear that the lowest frequency analyzed had to be raised from 4 Hz to 7 Hz to satisfy the criteria that (Cxy)2 ≥ 0.67.

Removed in the revised version: As long as muscle can be considered to be linear (a sinusoidal change in length produces a sinusoidal change in force), then system identification methods [78], [79] can be applied to extract a relationship between length x(t) and force y(t). We will give a brief overview of system identification methods here to make methods and results clearer. First, the time-domain signals (x(t) and y(t)) are transformed into an equivalent representation in the frequency-domain (X(s) and Y (s)) as a sum of scaled and shifted sine curves (Fig. 4B and C) using a Fourier transform [78]. In the frequency domain, we identify an LTI system of best fit H(s) that describes how muscle transforms changes in length into changes in force such that Y (s) = H(s) X(s). Next, we evaluate how H(s) scales the magnitude (gain) and shifts the timing (phase) of a sinusoid in X(s) into a sinusoid of the same frequency in Y (s) (Fig. 4D). This process is repeated across all frequency-matched pairs of input and output sinusoids to build a function of how muscle scales (Fig. 4E) and shifts (Fig. 4F) input length sinusoids into output force sinusoids. The resulting transformation turns two complicated time-domain signals (Fig. 4A) into a clear relationship in the frequency-domain that describes how muscle transforms length changes into force changes: a very slow (0Hz) length change will result in an output force that is scaled by 4.5 and is in phase (Fig. 4E and F), a 35 Hz sinusoidal length change will produce an output force that is scaled by 4.9 and leads the input signal by 24° (Fig. 4E and F), and frequencies between 0 Hz and 35 Hz will be between these two signals in terms of scaling and phase. These patterns of gain and phase can be used to identify a network of spring-dampers that is equivalent to the underlying linear system (the system in Fig. 4 A, E, and F is a 4.46 N/mm spring in parallel with a 0.0089 Ns/mm damper). Since experimental measurements often contain noise and small nonlinearities, the mathematical procedure used to estimate H(s) and the corresponding gain and phase profiles is more elaborate than we have described (see Appendix D for details).

Removed in the revised version: Kirsch et al. [5] used system identification methods to identify LTI mechanical systems that best describes how muscle transforms input length waveforms to output force waveforms. The resulting LTI system, however, is only accurate when the relationship between input and output is approximately linear. Kirsch et al. [5] used the coherence squared, (Cxy)2, between the input and output to evaluate the degree of linearity: Y (s) is a linear transformation of X(s) at frequencies in which (Cxy)2 is near one, while Y (s) cannot be described as a linear function of X(s) at frequencies in which (Cxy)2 approaches zero. By calculating (Cxy)2 between the length perturbation and force waveforms, Kirsch et al. [5] identified the bandwidth in which the muscle’s response is approximately linear. Kirsch et al. [5] set the lower frequency of this band to 4 Hz, and Fig. 3 of Kirsch et al. [5] suggests that this corresponds to (Cxy)2 ≥ 0.67 though the threshold for (Cxy)2 is not reported. The upper frequency of this band was set to the cutoff frequency of the low-pass filter applied to the input (15, 35, or 90 Hz). Within this bandwidth, Kirsch et al. [5] compared the response of the specimen to several candidate models and found that a parallel spring-damper fit the muscle’s frequency response best. Next, they evaluated the stiffness and damping coefficients that best fit the muscle’s frequency response [5]. Finally, Kirsch et al. evaluated how much of the muscle’s time-domain response was captured by the spring-damper of best fit by evaluating the variance-accounted-for (VAF) between the two time-domain signals Astonishingly, Kirsch et al. [5] found that a spring-damper of best fit has a VAF of between 78-99%12 when compared to the experimentally measured forces f EXP. By repeating this experiment over a variety of stimulation levels (using both electrical stimulation and the crossed-extension reflex) Kirsch et al. [5] showed that these stiffness and damping coefficients vary linearly with the active force developed by the muscle. Further, Kirsch et al. [5] repeated the experiment using perturbations that had a variety of lengths (0.4 mm, 0.8mm, and 1.6mm) and bandwidths (15Hz, 35Hz, and 90Hz) and observed a peculiar quality of muscle: the damping coefficient of best fit increases as the bandwidth of the perturbation decreases (See Figures 3 and 10 of Kirsch et al. [5] for details). Here we simulate Kirsch et al.’s experiment [5] to determine, first, the VAF of the VEXAT model and the Hill model in comparison to a spring-damper of best fit; second, to compare the gain and phase response of the models to biological muscle; and finally, to see if the spring-damper coefficients of best fit for both models increase with active force in a manner that is similar to the cat soleus that Kirsch et al. studied [5].

Removed in the revised version: To simulate the experiments of Kirsch et al. [5] we begin by creating the 9 stochastic perturbation waveforms used in the experiment that vary in perturbation amplitude (0.4mm, 0.8mm, and 1.6mm) and bandwidth (0-15 Hz, 0-35 Hz, and 0-90 Hz)13. The waveform is created using a vector that is composed of random numbers with a range of [−1, 1] that begins and ends with a series of zero-valued padding points. Next, a forward pass of a 2nd order Butterworth filter is applied to the waveform and finally the signal is scaled to the appropriate amplitude (Fig. 5). The muscle model is then activated until it develops a constant tension at a length of . The musculotendon unit is then simulated as the length is varied using the previously constructed waveforms while activation is held constant. To see how impedance varies with active force, we repeated these simulations at ten evenly spaced tensions from 2.5N to 11.5N. Ninety simulations are required to evaluate the nine different perturbation waveforms at each of the ten tension levels. The time-domain length perturbations and force responses of the modeled muscles are used to evaluate the coherence squared of the signal, gain response, and phase responses of the models in the frequency-domain. Since the response of the models might be more nonlinear than biological muscle, we select a bandwidth that meets (Cxy)2 > 0.67 but otherwise follows the bandwidths analyzed by Kirsch et al. [5] (see Appendix D for details).

Removed in the revised version: When coupled with an elastic-tendon, the 15 Hz perturbations show that neither model can match the VAF of Kirsch et al.’s analysis [5] (compare Fig. 6A to G), while at 90Hz the VEXAT model reaches a VAF of 89% (Fig. 6D) which is within the range of 78-99% reported by Kirsch et al. [5]. In contrast, the Hill model’s VAF at 90 Hz remains low at 58% (Fig. 6J). While the VEXAT model has a gain profile in the frequency-domain that closer to Kirch et al.’s data [5] than the Hill model (compare Fig. 6B to H and E to K), both models have a greater phase shift than Kirch et al.’s data [5] at low frequencies (compare Fig. 6C to I and F to L). The phase response of the VEXAT model to the 90 Hz perturbation (Fig. 6F) shows the consequences of Eqn. 16: at low frequencies the phase response of the VEXAT model is similar to that of the Hill model, while at higher frequencies the model’s response becomes similar to a spring-damper. This frequency dependent response is a consequence of the first term in Eqn. 16: the value of τ S causes the response of the model to be similar to a Hill model at lower frequencies and mimic a spring-damper at higher frequencies. Both models show the same perturbation-dependent phase-response, as the damping coefficient of best fit increases as the perturbation bandwidth decreases: compare the damping coefficient of best fit for the 15Hz and 90Hz profiles for the VEXAT model (listed on Fig. 6A. and D.) and the Hill model (listed on Fig. 6G. and J., respectively).

The closeness of each model’s response to the spring-damper of best fit changes when a rigid-tendon is used instead of an elastic-tendon. While the VEXAT model’s response to the 15 Hz and 90 Hz perturbations improves slightly (compare Fig. 6A-F to Fig. 16A-F in Appendix F), the response of the Hill model to the 15 Hz perturbation changes dramatically with the time-domain VAF rising from 55% to 85% (compare Fig. 6G-L to Fig. 16G-L in Appendix F). Although the Hill model’s VAF in response to the 15 Hz perturbation improved, the frequency response contains mixed results: the rigid-tendon Hill model’s gain response is better (Fig. 16H in Appendix F), while the phase response is worse in comparison to the elastic-tendon Hill model. While the rigid-tendon Hill model produces a better time-domain response to the 15 Hz perturbation than the elastic-tendon Hill model, this improvement has been made with a larger phase shift between force and length than biological muscle [5].

The gain and phase profiles of both models deviate from the spring-damper of best fit due to the presence of nonlinearities, even for small perturbations. Some of the VEXAT model’s nonlinearities in this experiment come from the tendon model (compare Fig. 6A-F to Fig. 16A-F in Appendix F), since the response of the VEXAT model with a rigid-tendon stays closer to the spring-damper of best fit. The Hill model’s nonlinearities originate from the underlying expressions for stiffness and damping of the Hill model, which are particularly nonlinear with a rigid-tendon model (Fig. 16G-L in Appendix F) The stiffness of a Hill model’s CE is heavily influenced by the partial derivative of which has a region of negative stiffness. Although is well approximated as being linear for small length changes, changes sign across . The damping of a Hill model’s CE also suffers from high degrees of nonlinearity for small perturbations about v M = 0 since the slope of is positive and large when shortening, and positive and small when lengthening (Fig. 2B). While Eqns. 26 and 27 are mathematically correct, the negative stiffness and wide ranging damping values predicted by these equations do not match experimental data [5]. In contrast, the stiffness and damping of the VEXAT’s CE do not change so drastically because these terms do not depend on the slope of the force-length relation, or the force-velocity relation (see Appendix B.4 for derivation).

By repeating the stochastic perturbation experiments across a range of isometric forces, Kirsch et al. [5] were able to show that the stiffness and damping of a muscle varies linearly with the active tension it develops (see Figure 12 of [5]). We have repeated our simulations of Kirsch et al.’s [5] experiments at ten nominal forces (spaced evenly between 2.5N and 11.5 N) and compared how the VEXAT model and the Hill model’s stiffness and damping coefficients compare to Figure 12 of Kirsch et al. [5] (Fig. 7). The stiffness and damping profile of the VEXAT model deviates a little from Kirsch et al.’s data [5] because XE’s dynamics at 35 Hz are still influenced by the Hill model embedded in Eqn. 16 (see Appendix B.4). Despite this, the VEXAT model develops similar stiffness and damping profile with either a viscoelastic-tendon (Fig. 7A & B) or a rigid-tendon (Fig. 7C & D). In contrast, when the Hill model is coupled with an elastic-tendon both its stiffness and damping are larger than Kirsch et al.’s data [5] at the higher tensions (Fig. 7A and B). This pattern changes when simulating a Hill model with a rigid-tendon: the model’s stiffness is slightly negative (Fig. 7C), while the model’s final damping coefficient is nearly three times the value measured by Kirsch et al. (Fig. 7D). Though a negative stiffness may seem surprising, Eqn. 26 shows a negative stiffness is possible at the nominal CE length of these simulations: just past the slope of the active force-length curve is negative and the slope of the passive force-length curve is negligible. The tendon model also affects the VAF of the Hill model to a large degree: the elastic-tendon Hill model has a low VAF 30-51% (Fig. 7A & B) while the rigid-tendon Hill model has a much higher VAF of 86%. Although the VAF of the rigid-tendon Hill model is acceptable, these forces are being generated in a completely different manner than those obtained from biological muscle, as Kirsch et al.’s data [5] indicate (Fig. 7C and D).

When coupled with an elastic-tendon, the stiffness (A.) and damping (B.) coefficients of best fit of both the VEXAT model and a Hill model increase with the tension developed by the MTU. However, both the stiffness and damping of the elastic-tendon Hill model are larger than Kirsch et al.’s coefficients (from Figure 12 of [5]), particularly at higher tensions. When coupled with rigid-tendon the stiffness (C.) and damping (D.) coefficients of the VEXAT model remain similar, as the values for and have been calculated to take the tendon model into account (see Appendix B.4 for details). In contrast, the stiffness and damping coefficients of the rigid-tendon Hill model differ dramatically from the elastic-tendon Hill model: while the elastic-tendon Hill model is too stiff and damped, the rigid-tendon Hill model is not stiff enough (compare A. to C.) and far too damped (compare B. to D.). Coupling the Hill model with a rigid-tendon increases the VAF from 30-51% to 86% but this improved accuracy is made using stiffness and damping that deviates from that of biological muscle [5].

Earlier version: When the VAF of the VEXAT and Hill model is evaluated across a range of nominal tensions (ten values from 2.5 to 11.5N), frequencies (15 Hz, 35 Hz, and 90 Hz), amplitudes (0.4mm, 0.8mm, and 1.6mm), and tendon types (rigid and elastic) two things are clear: first, that the VEXAT model’s 64-100% VAF is close to the 78-99% VAF reported by Kirsch et al. [5] while the Hill model’s 28-95% VAF differs (Fig. 8); and second, that there are systematic variations in VAF, stiffness, and damping across the different perturbation magnitudes and frequencies (see Tables 5 and 5 in Appendix E). Both models produce worse VAF values when coupled with an elastic-tendon (Fig. 8A, B, and C), though the Hill model is affected most: the mean VAF of the elastic-tendon Hill model is 67% lower than the mean VAF of the rigid-tendon model for the 0.4 mm 15 Hz pertur-bations (Fig. 8A). While the VEXAT model’s lowest VAF occurs in response to the low frequency perturbations (Fig. 8A) with both rigid and elastic-tendons, the Hill model’s lowest VAF varies with both tendon type and frequency: the rigid-tendon Hill model has its lowest VAF in response to the 1.6 mm 90 Hz perturbations (Fig. 8C) while the elastic-tendon Hill model has its lowest VAF in response to the 0.4 mm 15 Hz perturbations (Fig. 8A). It is unclear if biological muscle displays systematic shifts in VAF since Kirsch et al. [5] did not report the VAF of each trial.

Revised version: When the VAF of the VEXAT and Hill model is evaluated across a range of nominal tensions (ten values from 2.5 to 11.5N), frequencies (15 Hz, 35 Hz, and 90 Hz), amplitudes (0.4mm, 0.8mm, and 1.6mm), and tendon types (rigid and elastic) two things are clear: first, that the VEXAT model’s 64-100% VAF is close to the 78-99% VAF reported by Kirsch et al. [5] while the Hill model’s 28-95% VAF differs (Fig. 8); and second, that there are systematic vari-ations in VAF, stiffness, and damping across the different perturbation magnitudes and frequencies (see Tables 5 and 5 in Appendix E). Both models produce worse VAF values when coupled with an elastic-tendon (Fig. 8A, B, and C), though the Hill model is affected most: the mean VAF of the elastic-tendon Hill model is 67% lower than the mean VAF of the rigid-tendon model for the 0.4 mm 15 Hz perturbations (Fig. 8A). While the VEXAT model’s lowest VAF occurs in response to the low frequency perturbations (Fig. 8A) with both rigid and elastic-tendons, the Hill model’s lowest VAF varies with both tendon type and frequency: the rigid-tendon Hill model has its lowest VAF in response to the 1.6 mm 90 Hz perturbations (Fig. 8C) while the elastic-tendon Hill model has its lowest VAF in response to the 0.4 mm 15 Hz perturbations (Fig. 8A). It is unclear if biological muscle displays systematic shifts in VAF since Kirsch et al. [5] did not report the VAF of each trial.

Kirsch et al. [5] noted that the VAF of the spring-damper model of best fit captured between 78-99% across all experiments. We have repeated the perturbation experiments to evaluate the VAF across a range of conditions: two different tendon models, three perturbation bandwidths (15 Hz, 35 Hz, and 90 Hz), three perturbation magnitudes (0.4 mm, 0.8 mm and 1.6 mm), and ten nominal force values (spaced evenly between 2.5N and 11.5N). Each bar in the plot shows the mean VAF across all ten nominal force values, with the whiskers extending to the minimum and maximum value that occurred in each set. The mean VAF of the VEXAT model changes by up to 36% depending on the condition, with the lowest mean VAF occurring in response to the 1.6 mm 15 Hz perturbation with an elastic-tendon (A.), and the highest mean VAF occurring in response to the 90 Hz perturbations with the rigid-tendon (C.). In contrast, the mean VAF of the Hill model varies by up to 67% depending on the condition, with the lowest VAF occurring in the 15 Hz 0.4 mm trial with the elastic-tendon (A.), and the highest value VAF occurring in the 15 Hz 0.4 mm trial with the rigid-tendon (A.).

Herzog and Leonard [7] actively lengthened (A.) cat soleus muscles on the descending limb of the force-length curve (where in Fig. 2A) and measured the force response of the MTU (B.). After the initial transient at 2.4s the Hill model’s output force drops (B.) because of the small region of negative stiffness (C.) created by the force-length curve. In contrast, the VEXAT model develops steadily increasing forces between 2.4 − 3.4s and has a consistent level of stiffness (C.) and damping (D.).

In the VEXAT model we consider two different versions of the force-length relation for each titin segment (A): a linear extrapolation, and a WLC model extrapolation. Leonard et al. [8] observed that active myofibrils continue to develop increasing amounts of tension beyond actin-myosin overlap (B, grey lines with ±1 standard deviation shown). When this experiment is replicated using the VEXAT model (B., blue & magenta lines) and a Hill model (C. red lines), only the VEXAT model with the linear extrapolated titin model is able to replicate the experiment with the titin-actin bond slipping off of the actin filament at .

Active lengthening on the descending limb

Earlier version: We now turn our attention to the active lengthening insitu experiments of Herzog and Leonard [7]. During these experiments cat soleus muscles were actively lengthened by modest amounts (7-21% ) starting on the descending limb of the active-force-length curve in Fig. 2A). This starting point was chosen specifically because the stiffness of a Hill model may actually change sign and become negative because of the influence of the active-force-length curve on k M as shown in Eqn. 26 as ℓ M extends beyond . Herzog and Leonard’s [7] experiment is important for showing that biological muscle does not exhibit negative stiffness on the descending limb of the active-force-length curve. In addition, this experiment also highlights the slow recovery of the muscle’s force after stretching has ceased, and the phenomena of passive force enhancement after stimulation is removed. Here we will examine the 9 mm/s ramp experiment in detail because the simulations of the 3 mm/s and 27 mm/s ramp experiments produces similar stereotypical patterns (see Appendix G for details).

Revised version: We now turn our attention to the active lengthening insitu experiments of Herzog and Leonard [7]. During these experiments cat soleus muscles were actively lengthened by modest amounts starting on the descending limb of the active-force-length curve ( in Fig. 2A). This starting point was chosen specifically because the stiffness of a Hill model may actually change sign and become negative because of the influence of the active-force-length curve on k M as shown in Eqn. 26 as ℓ M extends beyond . Herzog and Leonard’s [7] experiment is important for showing that biological muscle does not exhibit negative stiffness on the descending limb of the active-force-length curve. In addition, this experiment also highlights the slow recovery of the muscle’s force after stretching has ceased, and the phenomena of passive force enhancement after stimulation is removed. Here we will examine the 9 mm/s ramp experiment in detail because the simulations of the 3 mm/s and 27 mm/s ramp experiments produces similar stereotypical patterns (see Appendix G for details).

When Herzog and Leonard’s [7] active ramp-lengthening (Fig. 9A) experiment is simulated, both models produce a force transient initially (Fig. 9B), but for different reasons. The VEXAT model’s transient is created when the lumped crossbridge spring (the term in Eqn. 15) is stretched. In contrast, the Hill model’s transient is produced, not by spring forces, but by damping produced by the force-velocity curve as shown in Eqn. 26.

After the initial force transient, the response of the two models diverges (Fig. 9B): the VEXAT model continues to develop increased tension as it is lengthened, while the Hill model’s tension drops before recovering. The VEXAT model’s continued increase in force is due to the titin model: when activated, a section of titin’s PEVK region remains approximately fixed to the actin element (Fig. 1C). As a result, the ℓ 2 element (composed of part of PEVK segment and the distal Ig segment) continues to stretch and generates higher forces than it would if the muscle were being passively stretched. While both the elastic and rigid-tendon versions of the VEXAT model produce the same stereotypical ramp-lengthening response (Fig. 9B), the rigid-tendon model develops slightly more tension because the strain of the MTU is solely borne by the CE.

Earlier version: In contrast, the Hill model develops less force during lengthening when it enters a small region of negative stiffness (Fig. 9B and C) because the passive-force-length curve is too compliant to compensate for the negative slope of the active force-length curve. Similarly, the damping coefficient of the Hill model drops substantially during lengthening (Fig. 9D). Equation 27 and Figure 2B shows the reason that damping drops during lengthening: d f V/dv M, the slope of the line in Fig. 2B, is quite large when the muscle is isometric and becomes quite small as the rate of lengthening increases.

Revised version: In contrast, the Hill model develops less force during lengthening when it enters a small region of negative stiffness (Fig. 9B and C) because the passive-force-length curve is too compliant to compensate for the negative slope of the active force-length curve. Similarly, the damping coefficient of the Hill model drops substantially during lengthening (Fig. 9D). Equation 27 and Figure 2B shows the reason that damping drops during lengthening: d f V/dv M, the slope of the line in Fig. 2B, is quite large when the muscle is iso-metric and becomes quite small as the rate of lengthening increases.

After the ramp stretch is completed (at time 3.4 seconds in Fig. 9B), the tension developed by the cat soleus recovers slowly, following a profile that looks strikingly like a first-order decay. The large damping coefficient acting between the titin-actin bond slows the force recovery of the VEXAT model. We have tuned the value of to for the elastic-tendon model, and for the rigid-tendon model, to match the rate of force decay of the cat soleus in Herzog and Leonard’s data [7]. The Hill model, in contrast, recovers to its isometric value quite rapidly. Since the Hill model’s force enhancement during lengthening is a function of the rate of lengthening, when the lengthening ceases, so too does the force enhancement.

Earlier version: Once activation is allowed to return to zero, Herzog and Leonard’s data shows that the cat soleus continues to develop a tension that is ΔfB above passive levels (Fig. 9B for t > 8.5s). The force ΔfB is known as passive force enhancement, and is suspected to be caused by titin binding to actin [76]. Since we model titin-actin forces using an activation-dependent damper, when activation goes to zero our titin model becomes unbound from actin. As such, both our model and a Hill model remain ΔfB below the experimental data of Herzog and Leonard (Fig. 9B) after lengthening and activation have ceased.

Revised version: Once activation is allowed to return to zero, Herzog and Leonard’s data shows that the cat soleus continues to develop a tension that is ΔfB above passive levels (Fig. 9B for t > 8.5s). The force ΔfB is known as passive force enhancement, and is suspected to be caused by titin bind-ing to actin [76]. Since we model titin-actin forces using an activation-dependent damper, when activation goes to zero our titin model becomes unbound from actin. As such, both our model and a Hill model remain ΔfB below the experimental data of Herzog and Leonard (Fig. 9B) after lengthening and activation have ceased.

Earlier version: Active lengthening beyond actin-myosin over-lap

Revised version: Active lengthening beyond actin-myosin overlap

Earlier version: One of the great challenges that remains is to decompose how much tension is developed by titin (Fig. 1C) separately from myosin (Fig. 1B) in an active sarcomere. Leonard et al.’s [8] active-lengthening experiment provides some insight into this force distribution problem because they recorded active forces both within and far beyond actin-myosin overlap. Leonard et al.’s [8] data shows that active force continues to develop linearly during lengthening, beyond actin-myosin overlap, until mechanical failure. When activated and lengthened, the myofibrils failed at a length of and force of , on average. In contrast, during passive lengthening myofibrils failed at a much shorter length of with a dramatically lower tension of of . To show that the extraordinary forces beyond actin-myosin overlap can be ascribed to titin, Leonard et al. [8] repeated the experiment but deleted titin using trypsin: the titin-deleted myofibrils failed at short lengths and insignificant stresses. Using the titin model of Eqn. 20 (Fig. 1A) as an interpretive lens, the huge forces developed during active lengthening would be created when titin is bound to actin leaving the distal segment of titin to take up all of the strain. Conversely, our titin model would produce lower forces during passive lengthening because the proximal Ig, PEVK, and distal Ig regions would all be lengthening together (Fig. 3A).

Revised version: One of the great challenges that remains is to decompose how much tension is developed by titin (Fig. 1C) separately from myosin (Fig. 1B) in an active sarcomere. Leonard et al.’s [8] active-lengthening experiment provides some insight into this force distribution problem because they recorded active forces both within and far beyond actin-myosin overlap. Leonard et al.’s [8] data shows that active force continues to develop linearly during lengthening, beyond actin-myosin overlap, until mechanical failure. When activated and lengthened, the myofibrils failed at a length of and force of , on average. In contrast, dur-ing passive lengthening myofibrils failed at a much shorter length of with a dramatically lower tension of of . To show that the extraordinary forces beyond actin-myosin overlap can be ascribed to titin, Leonard et al. [8] repeated the experiment but deleted titin using trypsin: the titin-deleted myofibrils failed at short lengths and insignificant stresses. Using the titin model of Eqn. 20 (Fig. 1A) as an interpretive lens, the huge forces developed during active lengthening would be created when titin is bound to actin leaving the distal segment of titin to take up all of the strain. Conversely, our titin model would produce lower forces during passive lengthening because the proximal Ig, PEVK, and distal Ig regions would all be lengthening together (Fig. 3A).

Since Leonard et al.’s experiment [8] was performed on skinned rabbit myofibrils and not on whole muscle, both the VEXAT and Hill models had to be adjusted prior to simulation (see Appendix H for parameter values). To simulate a rabbit myofibril we created a force-length curve [75] consistent with the filament lengths of rabbit skeletal muscle [60] (1.12µm actin, 1.63µm myosin, and 0.07µm z-line width) and fit the force-length relations of the two titin segments to be consistent with the structure measured by Prado et al. [58] of rabbit psoas titin consisting of a 70%-30% mix of a 3300kD and a 3400kD titin isoform (see Appendix B.3 for fitting details and Appendix H for parameter values). Since this is a simulation of a fibril, we used a rigid-tendon of zero length (equivalent to ignoring the tendon), and set the pennation angle to zero.

As mentioned in Sec. 2, because this experiment includes extreme lengths, we consider two different force-length relations for each segment of titin (Fig. 10A): a linear extrapolation, and an extension that follows the WLC model. While both versions of the titin model are identical up to , beyond the WLC model continues to develop increasingly large forces until all of the Ig domains and PEVK residues have been unfolded and the segments of titin reach a physical singularity: at this point the Ig domains and PEVK residues cannot be elongated any further without breaking molecular bonds (see Appendix B.3 for details). Our preliminary simulations indicated that the linear titin model’s titin-actin bond was not strong enough to support large tensions, and so we increased the value of from 71.9 to 975 (compare Tables 1 and 6 section H).

The Hill model was similarly modified, with the pennation angle set to zero and coupled with a rigid-tendon of zero length. Since the Hill model lacks an ECM element the passive-force-length curve was instead fitted to match the passive forces produced in Leonard et al.’s data [8]. No adjustments were made to the active elements of the Hill model.

When the slow active stretch (0.1µm/sarcomere/s) of Leonard et al.’s experiment is simulated [8] only the VEXAT model with the linear titin element can match the experimental data of Leonard et al. [8] (Fig. 10B). The Hill model cannot produce active force for lengths greater than since the active force-length curve goes to zero (Fig. 2A) and the model lacks any element capable of producing force beyond this length. In contrast, the linear titin model continues to develop active force until a length of is reached, at which point the titin-actin bond is pulled off the end of the actin filament and the active force is reduced to its passive value.

The WLC titin model is not able to reach the extreme lengths observed by Leonard et al. [8]. The distal segment of the WLC titin model approaches its contour length early in the simulation and ensures that the the titin-actin bond is dragged off the end of the actin filament at (Fig. 10B). After (Fig. 10B), the tension of the WLC titin model drops to its passive value but continues to increase until the contour lengths of all of the segments of titin are reached at . Comparing the response of the linear model to the WLC titin model two things are clear: the linear titin model more faithfully follows the data of Leonard et al. [8], but does so with titin segment lengths that exceed the maximum contour length expected for the isoform of titin in a rabbit myofibril.

This simulation has also uncovered a surprising fact: the myofibrils in Leonard et al.’s [8] experiments do not fail at , as would be expected by the WLC model of titin, but instead reach much greater lengths (Fig. 2B). Physically, it may be possible for a rabbit myofibril to reach these lengths (without exceeding the contour lengths of the proximal Ig, PEVK, and distal Ig segments) if the bond between the distal segment of titin and myosin breaks down. This would allow the large Ig segment, that is normally bound to myosin, to uncoil and continue to develop the forces observed by Leonard et al. [8]. Unfortunately the mechanism which allowed the samples in Leonard et al.’s experiments to develop tension beyond titin’s contour length remains unknown.

Force-length and force-velocity

Although the active portion of the Hill model is embedded in Eqn. 16, it is not clear if the VEXAT model can still replicate Hill’s force-velocity experiments [9] and Gordon et al.’s [10] force-length experiments. Here we simulate both of these experiments using the cat soleus model that we have used for the simulations described in Sec. 3.1 and compare the results to the force-length and force-velocity curves that are used in the Hill model and in Eqn. 16 of the VEXAT model.

Hill’s force-velocity experiment [9] is simulated by activating the model, and then by changing its length to follow a shortening ramp and a lengthening ramp. During short-ening experiments, the CE shortens from to with the measurement of active muscle force is made at Lengthening experiments are similarly made by measuring muscle force mid-way through a ramp stretch that begins at and ends at . When an elastic-tendon model is used, we carefully evaluate initial and terminal path lengths to accommodate for the stretch of the tendon so that the CE still shortens from to and lengthens from to .

Earlier version: The VEXAT model produces forces that differ slightly from the f V that is embedded in Eqn. 16 while the Hill model reproduces the curve (Fig. 11). The maximum shortening velocity of the VEXAT model is slightly weaker than the embedded curve due to the series viscoelasticity of the XE element. Although the model can be made to converge to the f V curve more rapidly by decreasing τ S this has the undesirable consequence of degrading the low-frequency response of the model during Kirsch et al.’s experiments [5] (particularly Fig. 6C., and F.). These small differences can be effectively removed by scaling by sV (Fig. 11A has sV = 0.95) to accommodate for the small decrease in force caused by the viscoelastic XE element.

Revised version: The VEXAT model produces forces that differ slightly from the f V that is embedded in Eqn. 16 while the Hill model reproduces the curve (Fig. 11). The maximum short-ening velocity of the VEXAT model is slightly weaker than the embedded curve due to the series viscoelasticity of the XE element. Although the model can be made to converge to the f V curve more rapidly by decreasing τ S this has the undesirable consequence of degrading the low-frequency response of the model during Kirsch et al.’s experiments [5] (particularly Fig. 6C., and F.). These small differences can be effectively removed by scaling by sV (Fig. 11A has sV = 0.95) to accommodate for the small decrease in force caused by the viscoelastic XE element.

When Hill’s [9] force-velocity experiment is simulated (A.), the VEXAT model produces a force-velocity profile (blue dots) that approaches zero more rapidly during shortening than the embedded profile f V(·) (red lines). By scaling by 0.95 the VEXAT model (magenta squares) is able to closely follow the force-velocity curve of the Hill model. While the force-velocity curves between the two models are similar, the time-domain force response of the two models differs substantially (B.). The rigid-tendon Hill model exhibits a sharp nonlinear change in force at the beginning (0.1s) and ending (0.21s) of the ramp stretch.

When Gordon et al.’s [10] passive and active force-length experiments are simulated, the VEXAT model (blue dots) and the Hill model (red lines) produce slightly different force-length curves (A.) and force responses in the time-domain (B.). The VEXAT model produces a right shifted active force-length curve, when compared to the Hill model due to the series elasticity of the XE element. By shifting the underlying curve by to the left the VEXAT model (magenta squares) can be made to exactly match the force-length characteristic of the Hill model.

Gordon et al.’s [10] force-length experiments were simulated by first passively lengthening the CE, and next by measuring the active force developed by the CE at a series of fixed lengths. Prior to activation, the passive CE was simulated for a brief period of time in a passive state to reduce any history effects due to the active titin element. To be consistent with Gordon et al.’s [10] experiment, we subtracted off the passive force from the active force before producing the active-force-length profile.

Earlier version: The simulation of Gordon et al.’s [10] experiment shows that the VEXAT model (Fig. 12A, blue dots) produces a force-length profile that is shifted to the right of the Hill model (Fig. 12A, red line) due to the series elasticity introduced by the XE. We can solve for the size of this right-wards shift by noting that Eqn. 16 will drive the to a length such that the isometric force developed by the XE is equal to that of the embedded Hill model allowing us to solve for the isometric strain of the XE. Since there are two viscoelastic XE elements per CE, the VEXAT model has an active force-length characteristic that shifted to the right of the embedded f L curve by a constant . This shift, ΔL, can be calibrated out of the model (Fig. 12 upper plot, magenta squares) by adjusting the f L(·) curve so that it is to the left of its normal position. Note that all simulations described in the previous sections made use of the VEXAT model with the calibrated force-length relation and the calibrated force-velocity relation.

Revised version: The simulation of Gordon et al.’s [10] experiment shows that the VEXAT model (Fig. 12A, blue dots) produces a force-length profile that is shifted to the right of the Hill model (Fig. 12A, red line) due to the series elasticity in-troduced by the XE. We can solve for the size of this right-wards shift by noting that Eqn. 16 will drive the to a length such that the isometric force developed by the XE is equal to that of the embedded Hill model allowing us to solve for the isometric strain of the XE. Since there are two viscoelastic XE elements per CE, the VEXAT model has an active force-length characteristic that shifted to the right of the embedded f L curve by a constant . This shift, ΔL, can be calibrated out of the model (Fig. 12 upper plot, magenta squares) by adjusting the f L(·) curve so that it is to the left of its normal position. Note that all simulations described in the previous sections made use of the VEXAT model with the calibrated force-length relation and the calibrated force-velocity relation.

Discussion

Earlier version: A muscle model is defined by the experiments it can replicate and the mechanisms it embodies. We have developed the VEXAT muscle model to replicate the force response of muscle to a wide variety of perturbations [5], [7], [8] while also retaining the ability to reproduce Hill’s force-velocity [9] experiment and Gordon et al.’s [10] force-length experiments. The model we have developed uses two mechanisms to capture the force response of muscle over a large variety of time and length scales: first, a viscoelastic crossbridge element that over brief time-scales appears as a spring-damper, and at longer time-scales mimics a Hill-model; second, a titin element that is capable of developing active force during large stretches.

Revised version: A muscle model is defined by the experiments it can replicate and the mechanisms it embodies. We have developed the VEXAT muscle model to replicate the force response of muscle to a wide variety of perturbations [5], [7], [8] while also retaining the ability to reproduce Hill’s force-velocity [9] experiment and Gordon et al.’s [10] force-length experiments. The model we have developed uses two mechanisms to capture the force response of muscle over a large variety of time and length scales: first, a viscoelastic crossbridge element that over brief time-scales appears as a springdamper, and at longer time-scales mimics a Hill-model; second, a titin element that is capable of developing active force during large stretches.

Earlier version: The viscoelastic crossbridge and titin elements we have developed introduce a number of assumptions into the model. While there is evidence that the activation-dependent stiffness of muscle originates primarily from the stiffness of the attached crossbridges [62], the origins of the activation-dependent damping observed by Kirsch et al. [5] have not yet been established. We assumed that, since the damping observed by Kirsch et al. [5] varies linearly with activation, the damping originates from the attached crossbridges. Whether this damping is intrinsic or is due to some other factor remains to be established. Next, we have also assumed that the force developed by the XE converges to a Hill model [17] given enough time (Eqn. 16). A recent experiment of Tomalka et al. [80] suggests the force developed by the XE might decrease during lengthening rather than increasing as is typical of a Hill model [17]. If Tomalka et al.’s [80] observations can be replicated, the VEXAT model will need to be adjusted so that the the XE element develops less force during active lengthening while the active-titin element develops more force. Finally, we have assumed that actin-myosin sliding acceleration (due to crossbridge cycling) occurs when there is a force imbalance between the external force applied to the XE and the internal force developed by the XE as shown in Eqn. 16. This assumption is a departure from previous models: Hill-type models [16], [17] assume that the tension applied to the muscle instantaneously affects the actin-myosin sliding velocity; Huxley models [11] assume that the actin-myosin sliding velocity directly affects the rate of attachment and detachment of crossbridges.

Revised version: The viscoelastic crossbridge and titin elements we have developed introduce a number of assumptions into the model. While there is evidence that the activationdependent stiffness of muscle originates primarily from the stiffness of the attached crossbridges [62], the origins of the activation-dependent damping observed by Kirsch et al. [5] have not yet been established. We assumed that, since the damping observed by Kirsch et al. [5] varies linearly with activation, the damping originates from the attached crossbridges. Whether this damping is intrinsic or is due to some other factor remains to be established. Next, we have also assumed that the force developed by the XE converges to a Hill model [17] given enough time (Eqn. 16). A recent experiment of Tomalka et al. [80] suggests the force developed by the XE might decrease during lengthening rather than increasing as is typical of a Hill model [17]. If Tomalka et al.’s [80] observations can be replicated, the VEXAT model will need to be adjusted so that the the XE element develops less force during active lengthening while the active-titin element develops more force. Finally, we have assumed that actin-myosin sliding acceleration (due to crossbridge cycling) occurs when there is a force imbalance between the external force applied to the XE and the internal force developed by the XE as shown in Eqn. 16. This assumption is a departure from previous models: Hill-type models [16], [17] assume that the tension applied to the muscle instantaneously affects the actin-myosin sliding velocity; Huxley models [11] assume that the actin-myosin sliding velocity directly affects the rate of attachment and detachment of crossbridges.

The active titin model that we have developed makes assumptions similar to Rode et al. [38] and Schappacher-Tilp et al. [40]: some parts of the PEVK segment bond to actin, and this bond cannot do any positive work on titin. The assumption that the bond between titin and actin cannot do positive work means that titin cannot be actively preloaded: it can only develop force when it is stretched. In contrast, two mechanisms have been proposed that make it possible for titin to be preloaded by crossbridge cycling: Nishikawa’s [39] winding filament theory and DuVall et al.’s [44] titin entanglement hypothesis. If titin were significantly preloaded by crossbridge cycling, the titin load path would support higher forces and the myosin-actin load path would bear less force. Accordingly, the overall stiffness of the CE would be reduced, affecting our simulations of Kirsch et al. [5]: lower myosin-actin loads mean fewer attached crossbridges, since crossbridges are stiff in comparison to titin, the stiffness of the CE would decrease (see Appendix A). Hopefully experimental work will clarify if titin can be actively preloaded by crossbridges in the future.

Earlier version: Both the viscoelastic crossbridge and active titin elements include simple myosin-actin and titin-actin bond models that improve accuracy but have limitations. First, the viscoelastic crossbridge element has been made to represent a population of crossbridges in which the contribution of any single crossbridge is negligible. Though it may be possible for the XE model to accurately simulate a maximally activated single sarcomere (which has roughly 20 attached crossbridges per half sarcomere [12], [81]) the accuracy of the model will degrade as the number of attached cross-bridges decreases. When only a single crossbridge remains, the XE model’s output will be inaccurate because it can only generate force continuously while a real crossbridge generates force discretely each time it attaches to, and detaches from, actin. Next, we have used two equations, Eqns. 16 and 21, that assume myosin-actin and titin-actin interactions are temperature-invariant and scale linearly with size ( and ). In contrast, myosin-actin interactions and some titin-actin interactions are temperature-sensitive [82], [83] and may not scale linearly with size. In Sec. 3.3 we had to adjust the active titin damping parameter, to simulate myofibril experiments [8], perhaps because the assumptions of temperature-invariance and size-linearity were not met: the initial value for came from fitting to in-situ experimental data [7] from whole muscle that was warmer (35 − 36.5° C vs 20 − 21° C) and larger ( of 42.9 mm vs. 10 − 15 μm) than the myofibrils [8]. While the cat soleus XE and titin model parameters (Table 1 G, H, and I) can be used as rough default values, these parameters should be refit to accurately simulate muscle that differs in scale or temperature from cat soleus. Finally, the VEXAT model in its current form ignores phenomena related to submaximal contractions: the shift in the peak of the force-length relation [84], and the scaling of the maximum shortening velocity [85]. We hope to include these phenomena in a later version of the VEXAT model to more accurately simulate submaximal contractions.

Revised version: Both the viscoelastic crossbridge and active titin elements include simple myosin-actin and titin-actin bond models that improve accuracy but have limitations. First, the viscoelastic crossbridge element has been made to represent a population of crossbridges in which the contribution of any single crossbridge is negligible. Though it may be possible for the XE model to accurately simulate a maximally activated single sarcomere (which has roughly 20 attached crossbridges per half sarcomere [12], [81]) the accuracy of the model will degrade as the number of attached cross-bridges decreases. When only a single crossbridge remains, the XE model’s output will be inaccurate because it can only generate force continuously while a real crossbridge generates force discretely each time it attaches to, and detaches from, actin. Next, we have used two equations, Eqns. 16 and 21, that assume myosin-actin and titin-actin interactions are temperature-invariant and scale linearly with size ( and ). In contrast, myosin-actin interactions and some titin-actin interactions are temperature-sensitive [82], [83] and may not scale linearly with size. In Sec. 3.3 we had to adjust the active titin damping parameter, to simulate myofibril experiments [8], perhaps because the assumptions of temperature-invariance and size-linearity were not met: the initial value for came from fitting to in-situ experimental data [7] from whole muscle that was warmer (35 − 36.5° C vs 20 − 21° C) and larger ( of 42.9 mm vs. 10 − 15 µm) than the myofibrils [8].

Added in the revised version: While the cat soleus XE and titin model parameters (Table 1 G, H, and I) can be used as rough default values, these parameters should be refit to accurately simulate muscle that differs in scale or temperature from cat soleus. Finally, the VEXAT model in its current form ignores phenomena related to submaximal contractions: the shift in the peak of the force-length relation [84], and the scaling of the maximum shortening velocity [85]. We hope to include these phenomena in a later version of the VEXAT model to more accurately simulate submaximal contractions.

The model we have proposed can replicate phenomena that occur at a wide variety of time and length scales: Kirsch et al.’s experiments [5] which occur over small time and length scales; and the active lengthening experiments of Herzog and Leonard [7] and Leonard et al. [8] which occur over physiological and supra-physiological length scales. In contrast, we have shown in Sec. 3.1 to 3.3 that a Hill-type model compares poorly to biological muscle when the same set of experiments are simulated. We expect that a Huxley model [11] is also likely to have difficulty reproducing Kirsch et al.’s experiment [5] because the model lacks an active damping element. Since titin was discovered [23] long after Huxley’s model was proposed [11], a Huxley model will be unable to replicate any experiment that is strongly influenced by titin such as Leonard et al.’s experiment [8].

Earlier version: Although there have been several more recent muscle model formulations proposed, none have the properties to simultaneously reproduce the experiments of Kirsch et al. [5], Herzog and Leonard [7], Leonard et al. [8], Hill [9], and Gordon et al. [10]. Linearized impedance models [14], [15] can reproduce Kirsch et al.’s experiments [5], but these models lack the nonlinear components needed to reproduce Gordon et al.’s force-length curve [10] and Hill’s force-velocity curve [9]. The models of Forcinito et al. [18], and Tahir et al. [41] have a structure that places a contractile element in series with an elastic-tendon. While this is a commonly used structure, at high frequencies the lack of damping in the tendon will drive the phase shift between length and force to approach zero. The measurements and model of Kirsch et al. [5], in contrast, indicate that the phase shift between length and force approaches ninety degrees with increasing frequencies. Though the Hill-type models of Haeufle et al. [20] and Gü nther et al. [22] have viscoelastic tendons, these models have no representation, of the viscoelasticity of the CE’s attached crossbridges. Similar to the Hill-type muscle model evaluated in this work [17], it is likely that models of Haeufle et al. [20] and Günther et al. [22] will not be able to match the frequency response of biological muscle. De Groote et al. [51], [52] introduced a short-range-stiffness element in parallel to a Hill model to capture the stiffness of biological muscle. While De Groote et al.’s [51], [52] formulation improves upon a Hill model it is unlikely to reproduce Kirsch et al.’s experiment [5] because we have shown in Sec. 3.1 that a Hill model has a frequency response that differs from biological muscle. Rode et al.’s [38] muscle model also uses a Hill model for the CE and so we expect that this model will have the same difficulties reproducing Kirsch et al.’s [5] experiment. Schappacher-Tilp et al.’s model [40] extends a Huxley model [11] by adding a detailed titin element. Similar to a Huxley model, Schappacher-Tilp et al.’s model [40] will likely have difficulty reproducing Kirsch et al.’s experiment [5] because it is missing an active damping element.

Revised version: Although there have been several more recent muscle model formulations proposed, none have the properties to simultaneously reproduce the experiments of Kirsch et al. [5], Herzog and Leonard [7], Leonard et al. [8], Hill [9], and Gordon et al. [10]. Linearized impedance models [14], [15] can reproduce Kirsch et al.’s experiments [5], but these models lack the nonlinear components needed to reproduce Gordon et al.’s force-length curve [10] and Hill’s forcevelocity curve [9]. The models of Forcinito et al. [18], and Tahir et al. [41] have a structure that places a contractile element in series with an elastic-tendon. While this is a commonly used structure, at high frequencies the lack of damping in the tendon will drive the phase shift between length and force to approach zero. The measurements and model of Kirsch et al. [5], in contrast, indicate that the phase shift between length and force approaches ninety degrees with increasing frequencies. Though the Hill-type models of Haeufle et al. [20] and Günther et al. [22] have viscoelastic tendons, these models have no representation, of the viscoelasticity of the CE’s attached crossbridges. Similar to the Hill-type muscle model evaluated in this work [17], it is likely that models of Haeufle et al. [20] and Günther et al. [22] will not be able to match the frequency response of biological muscle. De Groote et al. [51], [52] introduced a short-range-stiffness element in parallel to a Hill model to capture the stiffness of biological muscle. While De Groote et al.’s [51], [52] formulation improves upon a Hill model it is unlikely to reproduce Kirsch et al.’s experiment [5] because we have shown in Sec. 3.1 that a Hill model has a frequency response that differs from biological muscle. Rode et al.’s [38] muscle model also uses a Hill model for the CE and so we expect that this model will have the same difficulties reproducing Kirsch et al.’s [5] experiment. Schappacher-Tilp et al.’s model [40] extends a Huxley model [11] by adding a detailed titin element. Similar to a Huxley model, Schappacher-Tilp et al.’s model [40] will likely have difficulty reproducing Kirsch et al.’s experiment [5] because it is missing an active damping element.

Earlier version: While developing this model, we have come across open questions that we hope can be addressed in the future. How do muscle stiffness and damping change across the force-length curve? Does stiffness and damping change with velocity? What are the physical origins of the active damping observed by Kirsch et al. [5]? What are the conditions that affect passive-force enhancement, and its release? In addition to pursuing these questions, we hope that other researchers continue to contribute experiments that are amenable to simulation, and to develop musculotendon models that overcome the limitations of our model. To help others build upon our work, we have made the source code of the model and all simulations presented in this paper available online14.

Revised version: While developing this model, we have come across open questions that we hope can be addressed in the future. How do muscle stiffness and damping change across the forcelength curve? Does stiffness and damping change with velocity? What are the physical origins of the active damping observed by Kirsch et al. [5]? What are the conditions that affect passive-force enhancement, and its release? In addition to pursuing these questions, we hope that other researchers continue to contribute experiments that are amenable to simulation, and to develop musculotendon models that overcome the limitations of our model. To help others build upon our work, we have made the source code of the model and all simulations presented in this paper available online14.